```python
# Initialize Otter
import otter
grader = otter.Notebook("HW1.ipynb")
```

```python
import math
import pandas as pd
```

**Due Date: Sep 22, 2026 by 11:59 pm**

Each question carries 10 points. 

# Homework Category

Each homework assignment contains **auxiliary-skill** and **advanced-skill** tasks.

The assignment mark is the **highest mark for which both skill thresholds are satisfied**.

## Homework Assignment Marks

| Assignment Mark |                                Auxiliary Tasks |                                 Advanced Tasks |
| --------------- | ---------------------------------------------: | ---------------------------------------------: |
| **E**           |                               At least **90%** |                               At least **80%** |
| **M**           |                               At least **80%** |                               At least **70%** |
| **N**           |                               At least **70%** |                               At least **60%** |
| **X**           | Below the N threshold or insufficient evidence | Below the N threshold or insufficient evidence |

> **Important:** Both skill requirements must be satisfied. Therefore, the **lower demonstrated skill determines the assignment mark**.

For example, a score of **95% on auxiliary tasks** and **65% on advanced tasks** results in an **N** mark.

---

## Homework Flexibility and Reassessment

### One-Week Extension

A student may extend the deadline of **one homework assignment by seven calendar days without penalty**.

The student must notify the instructor **by the original deadline**.

Approved absences and accommodations do **not** use this extension.

### One Revision

A student may revise **one homework assignment** within **seven calendar days after feedback is provided**.

The revision must include:

* an explanation of the original mistake,
* a description of the changes made, and
* evidence of improved reasoning.

The new evidence **replaces the previous mark**.

> Extensions and revisions may not extend beyond the published final submission deadline.

---

## Homework Category Grade

The final Homework Category grade is determined by the student's overall profile across the seven homework assignments.

| Category Grade | Required Homework Profile                                                                                 |
| -------------- | --------------------------------------------------------------------------------------------------------- |
| **A**          | **E** on at least **6 of the 7** homework assignments, with the remaining assignment at **N or higher**   |
| **B**          | **M or higher** on at least **6 of the 7** homework assignments; the remaining assignment cannot be **X** |
| **C**          | **N or higher** on at least **5 of the 7** homework assignments                                           |
| **D**          | **N or higher** on at least **4 of the 7** homework assignments                                           |
| **F**          | The student does not satisfy the **D** requirement                                                        |


## Instructions: 

1. After completing **HW1**, download your completed `HW1.ipynb` file by selecting **File → Download** in JupyterLab.

2. Upload your completed `HW1.ipynb` file to the **HW1 Dropbox in D2L**.

3. This assignment is due by the **date and time listed in the course calendar**.

4. This is an **individual assignment**. As stated in the course syllabus, the use of AI tools is prohibited. You may not receive assistance from anyone other than the instructor, the TA, or the AARC. Any help received from the TA or AARC must be clearly documented.

5. If you use any resource other than the instructor or the course textbook/notes, you must document it in a **new Markdown cell in Question 20 of this notebook**. Include the source, what you learned from it, and specifically how you used that information in completing the assignment. Submissions that are copied, either partially or entirely, from another source will be subject to a significant penalty.

6. If you need help, visit the instructor during the posted office hours. If those times do not work with your schedule, email the instructor to arrange an alternative meeting time.

7. Passing the visible autograder tests does **not** guarantee that your answer is correct. The instructor may use additional hidden tests during grading. Full credit will be awarded only when your answers also pass these tests.

8. If you encounter any technical or assignment-related difficulties, please contact the instructor for assistance.


# Auxiliary Skills

### Question 1: 

Write a python program that converts a given minutes into the number of years and remaining days for the minutes. For simplicity, assume that a year has 365 days. 

For eg: 1000000000 (1 billion) minutes is approximately 1902 years and 214 days. 

*Hint*
- First convert the total minutes into days.

- Then use:

    - // to find how many complete 365-day years fit into those days.
    - % to find the remaining days after removing the complete years.

```python
mins = 1000000000 
MINS_IN_A_DAY = 24 * 60
DAYS_IN_A_YEAR = 365

total_days = ...
years = ...
remaining_days = ...

print("{} minutes is approximately {} years and {} days.".format(mins, years, remaining_days)) 
```

```python
grader.check("q1")
```

### Question 2: 

Write a program that stores any three-digit integer number between 100 and 999 and adds all the digits in the integer. For example, if an integer is 932, the sum of all its digits is 9 + 3 + 2 = 14.  

Your program should work for any digit and not just 932.

Do not use any loop to solve this problem.

*Hint*:
- Use % 10 to extract the last digit.
- Use // 10 to remove that last digit.

```python
number = 932

ones = ...
tens = ...
hundreds = ...

your_ans = ones + tens + hundreds
your_ans
```

```python
grader.check("q2")
```

### Question 3: 

Write a Python program that stores any positive integer and calculates the sum of all its digits.

For example:

Number: 932
Digit sum: 9 + 3 + 2 = 14

and:

Number: 52841
Digit sum: 5 + 2 + 8 + 4 + 1 = 20

- Your program should work for an integer containing **any** number of digits. You must use for loop to solve this question.
- Store your answer in the variable `your_ans`.

```python
number = 52841
total_digits = len(str(number)) 

your_ans = 0

...
```

```python
your_ans
```

```python
grader.check("q3")
```

### Question 4: 

You are given a list representing the following matrix:

```text
 1   2   3   1   2
-1   1   2   5  -3
 0   0   2   1   4
```

Assume that every row contains the same number of columns.

Write Python code to calculate the **sum of the elements in each row** and store the results in a list, `your_ans`.

For example, the sum of the second row is:

```text
-1 + 1 + 2 + 5 - 3 = 4
```

Therefore, after summing each row, then the final result should be:

```python
[9, 4, 7]
```

Your code should work for **any matrix with `m` rows and `n` columns**.

```python
my_list = [[1, 2, 3, 1, 2], [-1, 1, 2, 5, -3], [0, 0, 2, 1, 4]]
your_ans = []

for row in my_list: 
    ...
```

```python
your_ans
```

```python
grader.check("q4")
```

### Question 5: 

You are given a list representing the following matrix:

```text
 1   2   3   1   2
-1   1   2   5  -3
 0   0   2   1   4
```

Assume that every row contains the same number of columns.

Write Python code to calculate the **sum of the elements in each column** and store the results in a list.

For example, the sum of the second column is:

```text
2 + 1 + 0 = 3
```

Therefore, after summing each column, the final result should be:

```python
[0, 3, 7, 7, 3]
```

Your code should work for **any matrix with `m` rows and `n` columns**.


```python
my_list = [[1, 2, 3, 1, 2], [-1, 1, 2, 5, -3], [0, 0, 2, 1, 4]]
your_ans = []

...
```

```python
your_ans
```

```python
grader.check("q5")
```

# Advance Skills

### Question 6: 

You are given the following list of posts:

```python
posts = [
    "Python is amazing! #Python #DataScience #AI",
    "Learning #python and #AI is fun. #Coding",
    "I love #DataScience, #Python, and #MachineLearning",
    "#AI will change the world. #Python #ai",
    "Practicing #coding every day! #Python"
]
```

Write Python code that produces a dictionary, `your_ans`, containing each hashtag and the number of posts in which it appears.

Requirements:
- Hashtags are case-insensitive. For example, `#Python` and `#python` should be treated as the same hashtag.
- Remove punctuation such as commas, periods, and exclamation marks attached to hashtags.
- If the same hashtag appears more than once in the same post, count it only once for that post. This way, you will be counting the number of posts containing a hashtag.
- Store all hashtags in lowercase.

Expected output: 
```python
{
    "#python": 5,
    "#datascience": 2,
    "#ai": 3,
    "#coding": 2,
    "#machinelearning": 1
}
```

**Further Instructions:** 
- Your code should work for any posts. Your solution will be tested using additional posts that are not shown in the assignment.
- You may use only Python concepts covered in class so far: variables, expressions, loops, conditionals, strings/string methods, lists, tuples, and dictionaries.
- Do not use regular expressions, Counter, pandas, external libraries, or AI-generated code.
- After your code, create a markdown cell and briefly **answer** the following:
    - How does your program prevent the same hashtag from being counted twice within one post? <br>
    - What would need to change in your code if we wanted to count the total number of hashtag occurrences instead of the number of posts containing each hashtag? <br>
    - For this post, `#AI is useful! #Python, #AI`, trace your code by showing the value stored in each variable you created at every important step of the program. Then show how these values lead to the final result stored in the your_ans variable. <br>
- You should be prepared to explain or make a small modification to your solution after your submission of the HW assignment. The instructor may ask follow-up questions related to your solution during grading process.


*Hints*
1. Work with one post at a time
    - Use a for loop to process each string in posts.
    - For each post, think about creating a temporary list, `hashtags_in_this_post` that stores the unique hashtags found only in that post.

2. Break each post into words
    - A useful first step is: `words = post.split()`
    - Then inspect each word and decide whether it represents a hashtag.
    - Ask yourself:
        - How can you check whether a word starts with #?
        - Which string method can help? You may use `startswith()` method in string.
          
3. Clean the hashtag before storing it
    - Some hashtags may have punctuation attached, such as:
      ```text
      #Python,
      #AI!
      ```
    - Use string methods such as: `lower(), strip()` to clean the word before using it as a dictionary key.

4. Prevent duplicate counting within one post
    - For each post, maintain a temporary list.  Before adding a hashtag to this list, check whether it is already there.
    - This is especially important for a post such as: `#AI will change the world. #Python #ai`

5. Update the main dictionary only once per post
    - After you finish finding the unique hashtags in one post, loop through that temporary list and update `your_ans`.

```python
posts = [
    "Python is amazing! #Python #DataScience #AI",
    "Learning #python and #AI is fun. #Coding",
    "I love #DataScience, #Python, and #MachineLearning",
    "#AI will change the world. #Python #ai",
    "Practicing #coding every day! #Python"
]

your_ans = {}

...
...

```

```python
your_ans
```

```python
grader.check("q6")
```

### Question 7: 

The following dictionary stores student records:

```python
students = {
    "Alice": [("Python", 92), ("Statistics", 88), ("SQL", 95)],
    "Bob": [("Python", 78), ("Statistics", 82), ("SQL", 74)],
    "Charlie": [("Python", 95), ("Statistics", 91), ("SQL", 89)],
    "Tyler": [("Python", 95), ("Statistics", 91), ("SQL", 91)],
    "David": [("Python", 85), ("Statistics", 85), ("SQL", 85)]
}
```

Write Python code that creates a new dictionary, `your_ans`, where each student's name maps to a tuple containing:

`(average_score, highest_course, lowest_course)`

For example:
```python
{
     'Alice': (91.67, 'SQL', 'Statistics'),
     'Bob': (78.0, 'Statistics', 'SQL'),
     'Charlie': (91.67, 'Python', 'SQL'),
     'Tyler': (92.33, 'Python', 'Statistics'),
     'David': (85.0, 'Python', 'Python')
}
```

Requirements
- Round the average to 2 decimal places.
- Determine the course with the highest score.
- Determine the course with the lowest score.
- If multiple courses have the same highest or lowest score, use the course that appears first in the original list.
- Do not use max() or min() with a key= argument.

**Further Instructions**

Create a markdown cell to answer the following:

- For the student record `"David": [("Python", 85), ("Statistics", 85), ("SQL", 85)]`, trace your code by showing the values stored in each variable you created at every important step of the program. Clearly show how your code determines the average score, highest-scoring course, and lowest-scoring course, especially when multiple courses have the same score. Finally, show how the result for David is stored in the `your_ans` dictionary.

```python
students = {
    "Alice": [("Python", 92), ("Statistics", 88), ("SQL", 95)],
    "Bob": [("Python", 78), ("Statistics", 82), ("SQL", 74)],
    "Charlie": [("Python", 95), ("Statistics", 91), ("SQL", 89)],
    "Tyler": [("Python", 95), ("Statistics", 91), ("SQL", 91)],
    "David": [("Python", 85), ("Statistics", 85), ("SQL", 85)]
}
```

```python
your_ans = {}

...




your_ans
```

```python
grader.check("q7")
```

### Question 8: 

You are given a dictionary whose keys are tuples representing coordinates:

```python
encoded = {
    (2, 1): "T",
    (1, 3): "T",
    (3, 2): "N",
    (1, 1): "D",
    (2, 3): "A",
    (3, 1): "H",
    (1, 2): "A",
    (2, 2): "A",
    (3, 3): "K"
}
```

The keys have the form: `(row, column)`

Write Python code that reconstructs the hidden message by reading the characters:

- row by row,
- from the smallest row number to the largest,
- and within each row, from the smallest column number to the largest.

Your program should create a list of strings, `decoded_row_list`, where each string represents one decoded row.

For example:

```python
[
    "DAT",
    "TAA",
    "HNK"
]
```


Requirements
- Your program must work for dictionaries containing different numbers of rows and columns.
- Do not assume that there are exactly 3 rows or 3 columns.
- Do not hard-code row numbers, column numbers, characters, or the expected message.
- You may use only concepts covered in class so far: variables, loops, conditionals, strings, lists, tuples, dictionaries, and basic built-in functions.
- Do not use pandas, NumPy, or external libraries.
- Your solution will be tested using different encoded dictionaries that are not shown in the assignment.

**Further Instructions**

Create a markdown cell and answer the following questions:

(a) For the original encoded dictionary shown above, trace your program as it processes row 2.

Show:

    - The values of the important variables used to identify and process row 2.
    - The order in which the coordinates belonging to row 2 are processed.
    - How each character is added to the row string.
    - The value of the completed row string.
    - The contents of the decoded-row list immediately after row 2 has been processed.

(b) Suppose the dictionary contains the same coordinates and characters, but the entries are written in a completely different order.
Briefly explain why your program should still produce the same decoded message.

(c) Suppose the decoding rule changes so that:

    - rows are still read from smallest to largest;
    - but columns within each row must now be read from largest to smallest.

In 1–3 sentences, explain what part of your solution would need to change. Do not rewrite the entire program.

```python
encoded = {
    (2, 1): "T",
    (1, 3): "T",
    (3, 2): "N",
    (1, 1): "D",
    (2, 3): "A",
    (3, 1): "H",
    (1, 2): "A",
    (2, 2): "A",
    (3, 3): "K"
}
```

```python
rows = 3
cols = 3

decoded_row_list = [] 

...
```

```python
decoded_row_list
```

```python
grader.check("q8")
```

### Question 9: 

You are given two sentences:

```python
sentence1 = "Python makes data analysis simple and powerful"
sentence2 = "Data analysis with Python is powerful and useful"
```

Write Python code to determine which words occur in:

- both sentences,
- only the first sentence,
- only the second sentence.

Restrictions

- You are not allowed to use sets.

Your result should be a dictionary like:

```python
{
    "common": [...],
    "only_sentence1": [...],
    "only_sentence2": [...]
}
```

Requirements

- Comparison must be case-insensitive.
- Preserve the order in which words first appear.
- A word should appear only once in each result list.
- Do not use set().
- For the sentences above, "Python" and "python" should be considered the same word.


**Further Instructions**

Create a markdown cell to answer the following:

- Consider the following smaller example:
```python
sentence1 = "Python data Python AI"
sentence2 = "Data SQL python"
```
Trace your program and show the values stored in the important variables you created at each major step.

Your trace should clearly show:

    - The lists of words obtained from sentence1 and sentence2 after case normalization.
    - How your program determines whether each word belongs in common, only_sentence1, or only_sentence2.
    - How the three lists are stored in the final your_ans dictionary.


```python
sentence1 = "Python makes data analysis simple and powerful"
sentence2 = "Data analysis with Python is powerful and useful"
```

```python
lst1 = []

your_ans = {
                'common' : [], 
                'only_sentence1': [], 
                'only_sentence2': []
           }

...



```

```python
your_ans
```

```python
grader.check("q9")
```

### Question 10: 

You are given transaction data in the following form:

```python
transactions = [
    ("Alice", "Amazon", 35.50),
    ("Bob", "Walmart", 42.00),
    ("Alice", "Amazon", 35.50),
    ("Charlie", "Target", 75.00),
    ("Alice", "Target", 20.00),
    ("Bob", "Walmart", 42.00),
    ("Alice", "Amazon", 120.00),
    ("Charlie", "Amazon", 40.00),
    ("Alice", "Amazon", 35.50)
]
```

Each tuple contains:
`(customer, merchant, amount)`

Write Python code that creates the following dictionary, `your_ans`:
```python
{
    "Alice": {
        "total": ...,
        "largest": ...,
        "merchants": [...],
        "unique_transactions": [...],
        "duplicate_transactions": [...]
    },
    ...
}
```

For each customer:

`total` :  The total amount spent.

`largest` :  The largest single transaction amount.

`merchants` : A list of unique merchants visited by that customer, preserving their first appearance.

`unique_transactions`: A list of transactions that occured only once for that customer. 

`duplicate_transactions`: A list containing transactions that occurred more than once for that customer.

For example:
```python
("Amazon", 35.50)
```
is a duplicate transaction for Alice.

However, even though it appears three times, it should appear **only once** in the duplicate_transactions list.


```python
transactions = [
    ("Alice", "Amazon", 35.50),
    ("Bob", "Walmart", 42.00),
    ("Alice", "Amazon", 35.50),
    ("Charlie", "Target", 75.00),
    ("Alice", "Target", 20.00),
    ("Bob", "Walmart", 42.00),
    ("Alice", "Amazon", 120.00),
    ("Charlie", "Amazon", 40.00),
    ("Alice", "Amazon", 35.50)
]
```

```python
your_ans = {}

for item in transactions:
    name = item[0] 
    merchant = item[1]
    amount = item[2]

    ...

```

```python
your_ans
```

```python
grader.check("q10")
```

Tests for q1 (points: None)
```
>>> years == 1902 and remaining_days == 214
True
```


Tests for q10 (points: None)
```
>>> your_ans['Alice']['total'] == 246.5
True
```
```
>>> your_ans['Bob']['duplicate_transactions'][0] == ('Walmart', 42.0)
True
```
```
>>> your_ans['Charlie']['largest'] == 75.0
True
```
```
>>> your_ans['Charlie']['merchants'] == ['Target', 'Amazon']
True
```


Tests for q2 (points: None)
```
>>> your_ans == 14
True
```


Tests for q3 (points: None)
```
>>> your_ans == 20
True
```


Tests for q4 (points: None)
```
>>> your_ans == [9, 4, 7]
True
```


Tests for q5 (points: None)
```
>>> your_ans == [0, 3, 7, 7, 3]
True
```


Tests for q6 (points: None)
```
>>> your_ans['python'] == 5
True
```
```
>>> your_ans['ai'] == 3
True
```


Tests for q7 (points: None)
```
>>> your_ans['David'] == (85.0, 'Python', 'Python')
True
```


Tests for q8 (points: None)
```
>>> decoded_row_list[0] == 'DAT'
True
```
```
>>> decoded_row_list[-1] == 'HNK'
True
```


Tests for q9 (points: None)
```
>>> len(your_ans['common']) == 5
True
```
```
>>> len(your_ans['only_sentence1']) == 2
True
```
```
>>> len(your_ans['only_sentence2']) == 3
True
```
