```python
# Initialize Otter
import otter
grader = otter.Notebook("Numpy_Part2.ipynb")
```

```python
import numpy as np
```

# Numpy Part 2

In Part 1, we learned what NumPy arrays are, why they are useful, and how to perform familiar operations such as indexing and slicing. In this lesson, we will explore one of NumPy’s most powerful features: performing computations between an array and a scalar value, as well as element-wise operations between two arrays. We will also learn how to use several NumPy functions and methods to perform common mathematical computations.

### Scalar computation on a numpy

A **scalar** is one value, such as `2`, `-5`, or `0.25`. When a scalar and a NumPy array are used in an arithmetic expression, NumPy applies the operation to **every element**.

This behavior is a simple form of **broadcasting** (we will learn in our next lesson). NumPy conceptually stretches the scalar to match the array.

| Expression | Calculation performed |
|---|---|
| `values + 3` | Add 3 to every element |
| `values - 3` | Subtract 3 from every element |
| `values * 3` | Multiply every element by 3 |
| `values / 3` | Divide every element by 3 |
| `values ** 2` | Square every element |

<img src="pics/broadcast1.jpeg" width=500px />

```python
arr = np.array([1, 2, 3])
arr * 2
```

```python
values = np.array([1, 2, 3, 4])

print("Original:  ", values)
print("Times 2:   ", values * 2)
print("Plus 10:   ", values + 10)
print("Squared:   ", values ** 2)
print("Divided by 2:", values / 2)

```

<div class="np-card np-tip">
<h4>Why this matters</h4>

The expression `values * 2` replaces a loop such as “visit each position and multiply its value by 2.” This style is called a **vectorized operation**. It is concise and is usually faster than a Python loop for large arrays.
</div>

The original array is unchanged unless we assign a result back to it or another variable.

### Computation between two numpy arrays

When two arrays have the same shape, arithmetic is performed **element by element**. The value at index `0` is combined with the value at index `0`, the value at index `1` with the value at index `1`, and so on.

```python
a = np.array([10, 20, 30])
b = np.array([1, 2, 3])

print("a + b =", a + b)
print("a - b =", a - b)
print("a * b =", a * b)
print("a / b =", a / b)
```

### Example: changes in graduation totals

The arrays below store graduation totals for the same four colleges in the same order.

| Index | College | 2019 | 2022 |
|:---:|---|---:|---:|
| 0 | Science and Mathematics | 457 | 532 |
| 1 | Arts | 385 | 433 |
| 2 | Business | 573 | 501 |
| 3 | Education | 382 | 293 |

Create a Python code to determine the difference in the number of students graduating from each college in 2019 compared to 2022.To calculate the change from 2019 to 2022, use:

$$\text{change}=\text{new value}-\text{old value}$$

A positive result means an increase; a negative result means a decrease.

```python
studs_in_2019 = np.array([457, 385, 573, 382])
studs_in_2022 = np.array([532, 433, 501, 293])

diff = studs_in_2022 - studs_in_2019
diff 
```

**Keep in mind that for two NumPy arrays to be computed together, they must have the same shape.** 

In this example, both `studs_in_2019` and `studs_in_2022` have a shape of (4,).

```python
np.array([1, 2, 3]) + np.array([10, 20])  # ValueError
```

Always inspect `.shape` when an array calculation behaves unexpectedly. Shapes `(3,)` and `(2,)` are not compatible: their shapes are different. NumPy would raise a `ValueError`.
</div>

# Numpy methods 


## Creating sequences

### `np.arange()`

`np.arange()` creates evenly spaced values using a familiar start–stop–step pattern:

```python
np.arange(start, stop, step)
```

- `start` is included.
- `stop` is excluded.
- `step` controls the distance between consecutive values.
- With one argument, NumPy uses `start=0` and `step=1`.


```python
print("One argument:   ", np.arange(5))
print("Start and stop:", np.arange(3, 9))
print("With a step:   ", np.arange(3, 30, 5))
print("Counting down: ", np.arange(5, 0, -1))
```

The stop value is excluded. For example, `np.arange(3, 9)` ends at `8`, and `np.arange(5, 0, -1)` ends at `1`.

```python
np.arange(10) # here start = 0 and step = 1 are set by default
```

```python
np.arange(3, 9) # here start = 3, end = 8, step = 1 by default
```

```python
np.arange(3, 30, 5) # start = 3, end = 29, step = 5 
```

```python
np.arange(1.5, -2, -0.5) # start = 1.5, end = -1.99999.... , step = -0.5
```

Numpy module provides us many methods that can be applied on numpy array. 

### Summarizing an array

Functions that combine many values into one summary are often called **reductions**.

<div class="np-grid">
  <div class="np-mini"><strong><code>np.sum(a)</code></strong><br>Total</div>
  <div class="np-mini"><strong><code>np.mean(a)</code></strong><br>Arithmetic mean</div>
  <div class="np-mini"><strong><code>np.min(a)</code></strong><br>Smallest value</div>
  <div class="np-mini"><strong><code>np.max(a)</code></strong><br>Largest value</div>
  <div class="np-mini"><strong><code>np.prod(a)</code></strong><br>Product of all values</div>
</div>


```python
temperatures = np.array([71, 74, 79, 76, 72])

print("Total:  ", np.sum(temperatures))
print("Mean:   ", np.mean(temperatures))
print("Minimum:", np.min(temperatures))
print("Maximum:", np.max(temperatures))
print("Product of [2, 3, 4]:", np.prod(np.array([2, 3, 4])))

```

```python
arr = np.array([1, 2, 3])
arr = np.arange(5)
np.sum(arr) # add all elements of an array
```

```python
np.mean(arr) # returns the mean of all the elements of the array
```

```python
np.prod(arr) # returns the product of all the elements of the array
```

## Boolean array

When all the elements of a numpy array contains `True` or `False` values, then it's a boolean array.

Three functions are especially useful with Boolean arrays:

- `np.all(condition)` is `True` only if **every** element is `True`.
- `np.any(condition)` is `True` if **at least one** element is `True`.
- `np.count_nonzero(condition)` counts the `True` values.


```python
scores = np.array([88, 67, 91, 73, 59])
passed = scores >= 70

print("Scores: ", scores)
print("Passed? ", passed)
```

```python
print("Did every student pass?", np.all(passed))
print("Did any student pass?  ", np.any(passed))
print("Number who passed:     ", np.count_nonzero(passed))
```

These functions also accept numerical arrays. NumPy treats zero as `False` and a nonzero value as `True`.

```python
arr = np.array([0, 5, 3, 0])
np.all(arr) # returns True if all elements are True or non-zero values
```

```python
barr = np.array([True, True, True])
np.all(barr)
```

```python
barr = np.array([True, True, True, False])
np.all(barr)
```

```python
arr = np.array([0, 5, 3, 0])
np.any(arr) # returns True if any of the elements are True or non-zero value 
```

```python
barr = np.array([False, False, False])
np.any(barr)
```

```python
barr = np.array([False, False, False, True])
np.any(barr)
```

```python
arr = np.array([0, 5, 3, 0, 8])
np.count_nonzero(arr) # counts the number of True or non-zero elements in the array
```

## Transforming arrays

The following functions return new arrays rather than reducing the data to one number.

### Consecutive differences: `np.diff()`

`np.diff()` subtracts each element from the element immediately after it:

$$[x_0,x_1,x_2,x_3] \longrightarrow [x_1-x_0,\;x_2-x_1,\;x_3-x_2]$$

Because it measures gaps between adjacent values, the result has one fewer element.


```python
np.diff(arr) # returns the difference of the consecutive elements 
```

### Rounding: `np.round()`

Use the second argument to choose the number of digits after the decimal point.


```python
arr = np.array([0.5, 0.239, 0.8238])
np.round(arr, 2) # returns the round of each number to the nearest integer
```

### Running totals and products

- `np.cumsum()` produces a running total.
- `np.cumprod()` produces a running product.

```python
arr = np.array([1, 2, 3, 4])
np.cumprod(arr) # a cumulative product
```

```python
np.cumsum(arr) # a cumulative sum
```

### Common mathematical transformations

| Function | Meaning | Input caution |
|---|---|---|
| `np.exp(a)` | $e^x$ for each value $x$ | Can grow very quickly |
| `np.log(a)` | Natural logarithm, $\ln(x)$ | Use positive inputs |
| `np.sqrt(a)` | Square root | Use nonnegative inputs for real results |

```python
np.exp(arr) # exponentiate each element 
```

```python
np.log(arr) # natural logarithm of each element
```

```python
np.sqrt(arr) # square root of each element
```

### Sorting: `np.sort()`

`np.sort(a)` returns a sorted copy; it does not change `a`.


```python
np.sort(arr) # sort elements in ascending order
```

### Finding the position of an extreme value

`np.argmax()` returns the index of the first maximum value, while `np.argmin()` returns the index of the first minimum value.


```python
temperatures = np.array([71, 74, 79, 76, 72])

hottest_index = np.argmax(temperatures)
coldest_index = np.argmin(temperatures)

print("Index of first maximum:", hottest_index)
print("Maximum temperature:    ", temperatures[hottest_index])
print("Index of first minimum:", coldest_index)
print("Minimum temperature:    ", temperatures[coldest_index])
```

## Lesson summary

| Goal | Main tool or idea |
|---|---|
| Apply one calculation to all values | Array–scalar arithmetic |
| Pair corresponding values | Elementwise array arithmetic |
| Combine compatible shapes | Broadcasting |
| Generate a sequence with a step | `np.arange()` |
| Produce one summary value | `sum`, `mean`, `min`, `max`, `prod` |
| Locate an extreme value | `argmin`, `argmax` |
| Test all, any, or count conditions | `all`, `any`, `count_nonzero` |
| Measure changes | `diff` |
| Build running results | `cumsum`, `cumprod` |
| Transform values | `round`, `exp`, `log`, `sqrt` |
| Sort without changing the original | `np.sort()` |

<div class="np-card np-tip">
<h4>One habit to keep</h4>

When a result surprises you, inspect the array's **values**, **shape**, and **data type**:

```python
print(arr)
print(arr.shape)
print(arr.dtype)
```

These three checks solve many beginner NumPy problems.
</div>


# KC4_Part 2

Answer the following questions and upload this notebook to **KC4_Part2 dropbox**. The due date is **Sep 23 by 11:59 pm.**

Do not use `for` loops to answer these questions. Try to use numpy concepts to solve them.

<!-- BEGIN QUESTION -->

### Question 1

What is the output of the following python code?  Please determine this without executing any code.

```python
arr = np.array([1, 2, 3]) + 2
arr
```

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

### Question 2

What is the output of the following python code?  Please determine this without executing any code.

```python
arr = np.array([1, 2, 3]) < 3
arr
```

_Type your answer here, replacing this text._

<!-- END QUESTION -->

### Question 3

Given the temperatures of five cities in Celsius—14, 53, 19, and 25—write a Python program to convert these temperatures to Fahrenheit. The formula for converting Celsius to Fahrenheit is:

$temp\_in\_fahr = \frac{9 * temp\_in\_celcius}{5} + 32$

Store your answer in a variable, `temps_in_fahr`. 

Hint: 
1. Create a numpy array, which stores the temperatures in celcius
2. Then use the scalar computation trick to compute the temperatures in fahrenheit all at once

```python
temps_in_cel = np.array(...)
temps_in_cel = ...
temps_in_fahr = ...
temps_in_fahr 
```

```python
grader.check("q3")
```

### Question 4

Suppose you are given two matrices A and B, where  $A = \begin{bmatrix}
1 & 2 \\
1 & 1 \\
10 & 20 \\
\end{bmatrix}$ and $B = \begin{bmatrix}
1 & 1 \\
2 & 2 \\
3 & 3 \\
\end{bmatrix}$ . Write a python code to add these two matrices together and store your answer in a variable, `C`.

```python
A = ...
B = ...
C = ...
C
```

```python
grader.check("q4")
```

### Question 5

Write a python code to find the number of positive integers in the matrix $A= \begin{bmatrix}
1 & -2 & -3 & 4 \\
1 & 1 & -1 & -1 \\
2 & -2 & -1 & 0 \\
\end{bmatrix}$

*Hint*: First, using scalar computation to convert the boolean matrix into True/False, based on the number is positive and negative. Second, look into `np.sum()` or `np.count_nonzero()` to count the number of True values in the boolean matrix


```python
A = np.array([[1, -2, -3, 4], [1, 1, -1, -1], [2, -2, -1, 0]]) 
your_ans = ...
your_ans
```

```python
grader.check("q5")
```

<!-- BEGIN QUESTION -->

For each of the Python expressions below, write the output when the expression is evaluated. If the
expression evaluates to an array, you should format your answer like so: `array([..., ..., ...])`. If the expression fails to run, then write `Error` as your answer. You
may assume the standard import:

```import numpy as np```  

Do not run any code to answer these questions.  You may run it later to check your answer.

### Question 6

```python
np.array([5, 6, -3]) == 4
```

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

### Question 7


```python
np.sum(np.array([5, 2, 0]) < 3)
```

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

### Question 8

```python
np.array([1, 2, 3]) - np.array([1, 1, 1])
```

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

### Question 9

```python
np.array([4, 4, 5]) - np.array([2, 2])
```

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

### Question 10

```python
np.arange(1, 10, 2)
```

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

### Question 11

```python
arr = np.arange(10, 100, 10)
new_arr = arr[np.arange(1, 5)]
print(new_arr)
```

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

### Question 12
```python
arr = np.arange(10, 100, 10)
new_arr = arr[np.arange(1, 7, 2)]
print(new_arr)
```

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

### Question 13
```python
arr = np.arange(10, 100, 10)
new_arr = arr[np.arange(1, 10, 2)]
print(new_arr)
```

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

### Question 14

```python
np.cumsum(np.array([1, 2, 4]))
```

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

### Question 15

#### Estimate $\pi$ using Leibniz's formula

Leibniz has given us formula to estimate the value of $\pi$ using the following equation: 
$$\pi = 4 * \left(1 - \frac{1}{3} +  \frac{1}{5} -  \frac{1}{7} + \frac{1}{9} -  \frac{1}{11} + ... \right)$$

We are going to verify if this equation is true by computing the right hand side of equation, if the result is approximately 3.14 (close to the value of $\pi$) then the equation is correct. 

Instead of finding the sum infinitely, we will caculate it upto large value (say 100001), then the equation becomes:
$$\left(1 - \frac{1}{3} + \frac{1}{5} - \frac{1}{7} + \frac{1}{9} - \frac{1}{11} + ... - \frac{1}{99999} + \frac{1}{100001} \right)$$

Now we will separate this equation by grouping the negative fractions and positive fractions together respectively. So the equation becomes:
$$\left(1 + \frac{1}{5} + \frac{1}{9} + ... + \frac{1}{100001} \right) - \left(\frac{1}{3} + \frac{1}{7} + \frac{1}{11} + ... + \frac{1}{99999} \right)$$

Therefore, let's first calculate the positive fractions first. 
$$\left(1 + \frac{1}{5} + \frac{1}{9} + ... + \frac{1}{100001} \right)$$

We see pattern in the denominator 1, 5, 9, ..., 100001. 
Therefore, we can store these denominators in a numpy array. 

Then, we can divide them by 1, and then add them together. That will be the output for the sum of all positive fractions. 

```python
left_equation = ...


right_equation = ...

pi = 4 * (left_equation - right_equation)
pi
```

```python
grader.check("q15")
```

<!-- END QUESTION -->



Tests for q15 (points: None)
```
>>> np.isclose(3.14, np.round(pi, 2)).item()
True
```


Tests for q3 (points: None)
```
>>> temps_in_fahr[-1] == 77
True
```
```
>>> len(temps_in_fahr) == 4
True
```


Tests for q4 (points: None)
```
>>> C.shape == (3, 2)
True
```
```
>>> C[-1, -1].item() == 23
True
```


Tests for q5 (points: None)
```
>>> your_ans == 5
True
```
