```python
# Initialize Otter
import otter
grader = otter.Notebook("DataStructures.ipynb")
```

<div style="
    background: linear-gradient(135deg, #3b1d6b 0%, #6d28d9 55%, #8b5cf6 100%);
    color: white;
    padding: 36px 40px;
    border-radius: 18px;
    margin: 8px 0 22px 0;
    box-shadow: 0 10px 28px rgba(55, 30, 100, 0.18);
">
    <div style="
        display:inline-block;
        background:rgba(255,255,255,0.16);
        padding:6px 13px;
        border-radius:18px;
        font-size:12px;
        font-weight:700;
        letter-spacing:0.6px;
        margin-bottom:12px;
    ">
        INTRO TO DATA ANALYTICS • PYTHON FOUNDATIONS
    </div>
    <h1 style="margin:0; font-size:39px; line-height:1.15; color:white;">
        🧱 Python Data Structures
    </h1>
    <p style="margin:12px 0 0 0; font-size:17px; line-height:1.55; max-width:900px;">
        Learn how Python organizes text and collections using strings, lists, tuples, and dictionaries,
        and practice the operations that appear most often in introductory data analytics.
    </p>
</div>

<div style="
    border-left: 5px solid #7c3aed;
    background: #f7f4ff;
    padding: 18px 22px;
    border-radius: 12px;
    margin: 10px 0 22px 0;
">
<h3 style="margin-top:0;">🎯 Learning Objectives</h3>

By the end of this notebook, you should be able to:

- Explain the purpose of Python's core data structures.
- Use **strings** with indexing, slicing, membership tests, and common string methods.
- Create, access, modify, add, remove, and copy items in **lists**.
- Explain why lists are **mutable** while strings and tuples are **immutable**.
- Create and unpack **tuples**.
- Create, access, modify, add, and remove **dictionary** entries.
- Work with **nested lists and dictionaries**.
- Choose an appropriate data structure for a simple data analytics task.

</div>

<div style="
    background:#eef6ff;
    border:1px solid #cfe4ff;
    padding:16px 20px;
    border-radius:12px;
    margin-bottom:24px;
">
<strong>💡 Before you begin</strong><br><br>

Python data structures let us organize information in ways that make it easier to access and manipulate.

| Structure | Example | Ordered? | Mutable? | Main idea |
|---|---|---:|---:|---|
| **String** | `"Python"` | Yes | No | Sequence of characters |
| **List** | `[10, 20, 30]` | Yes | Yes | Editable sequence |
| **Tuple** | `(10, 20, 30)` | Yes | No | Fixed sequence |
| **Dictionary** | `{"name": "Ana"}` | Yes* | Yes | Key → value mapping |

<small>*Dictionaries preserve insertion order in modern Python, but their defining feature is key-based lookup.</small>
</div>

## 1. Strings

A **string** is a sequence of zero or more characters in Python. Strings are surrounded by double quotation marks or single quotation marks. 
For eg:

```python
print("Hello World")
print("") # this is empty string which has 0 character
```

### Indexing in String

Each character in a string maintains an index starting from 0 (shown in the table below):

|Index --> | 0 | 1 | 2 | 3 | 4 |
|----| ---- | --- | --- | ---- | ---- |
| | H | e | l | l | o |

We can use indexing to access any character of a string. 

For eg: To access the first character of the string, we can do: 

```python
my_str = "Hello"
my_str[0] # get character at position index 0
```

### String length

To get the length of a string, we can use the in-built python function, `len`

```python
len(my_str) # Hello contains total 5 characters
```

### Check String

To check if a substring is present in a string, we can use `in` keyword. 

```python
'o' in my_str
```

```python
my_name = "Dipak Singh"
'Sing' in my_name
```

### Slicing String

To extract a substring from a string, you can use slicing. Slicing involves specifying the starting and ending index positions within brackets, where the ending index is exclusive. For example, consider the following sentence:
|Index --> | 0 | 1 | 2 | 3 | 4 | 5 | 6 | 7 | 8 | 9 | 10 | 11 | 12 |  13 |14 | 15 | 16 | 17 | 18 | 19 | 20 |
| --- |--- | ---| ---|--- |--- | ---| ---| ---| ---| ---| ---| ---| ---| ---| ---| ---| ---|--- |--- |--- |--- |
| | I | | l | o | v | e |  | D | a | t | a | | A | n | a | l | y | t | i | c | s |

```python
my_var = "I love Data Analytics"
my_var[2:6] # the index number 6 is exclusive here
```

```python
my_var[:5] # By leaving out the start index, the range will start at the first character
```

```python
my_var[12:] # By leaving out the last index, the range will end at the last character
```

```python
# You can even use negative indexing where the index starts from the last character with index position -1 

my_var[-1]
```

```python
my_var[-6]
```

You can even slice characters of a string by adding a step size to it. For eg: 

```python
my_var[::2]
```

Consider the string, 'dog'. What do you think, will be the output of the following Python code? Explain why. 

```python
my_str = "dog"
print(my_str[-1::-1])
```

### String Methods

Python provides many useful **string methods** that can be called on string objects. A complete list of such functions can be found [here](https://docs.python.org/3/library/stdtypes.html#string-methods)

#### Change string to upper/lower case

The `upper` method is used to convert a string to upper case. 

```python
my_var = "Hello World"
my_var.upper() # Notice how the method is invoked by using the string followed by a dot.
```

```python
my_var.lower() # similarly, you can use lower to change to lower case
```

### Count String

The `count` method is used to count the number of times a substring appears in a string. 

```python
tongue_twister = "I scream, you scream, we all scream for ice cream."
tongue_twister.count('cream')
```

#### Remove whitespaces from beginning or end of a string

The `strip` method can be used to remove whitespaces at the start or end of a string.

```python
my_var = "    Hello         World            "
my_var.strip() # Notice how it only removes the spaces at the start and the end, but not in the middle
```

```python
my_var = "        Hello World!!!!!!!!!"
my_var.strip('! ') # You can specify multiple characters as a string argument to remove them from the beginning and end of the string.
```

#### Replace String

The `replace` method replaces a string with another string

```python
my_var = "I love CSCI1462"
my_var.replace("CSCI1462", "Data Analytics")
```

#### Split String

We can use the `split` method to divide a string into multiple substrings based on a delimiter. This method returns a list containing the resulting substrings.

```python
fruits = "Apple, Banana, Orange, Kiwi"
fruits.split(", ") # Splitting strings based on , character. Note the space after the comma
```

#### String Concatenation

You can use the `+` operator to combine two strings together. 

```python
s1 = "I love"
s2 = "Data Analytics"
print(s1 + " " + s2)
```

#### Escape Character within a String

If you want to add a special character within a string, then you may add them by using an escape character. 

For eg: Suppose we want to print the string, `He said, "How are you?"`

my_var = "He said, "How are you?"" 
print(my_var)

Oops! It seems the Python interpreter is having trouble distinguishing the double quotation marks within the string from the ones that delimit the string itself. To resolve this, you should use the escape character, a backslash (\), to indicate that the double quotation marks inside the string are meant to be part of the string rather than delimiters. Here’s how you can do it:

```python
my_var = "He said, \"How are you?\""
print(my_var)
```

For more detailed list of other String methods, click [here](https://www.w3schools.com/python/python_ref_string.asp). 

<div style="
    border-left:5px solid #f59e0b;
    background:#fffbeb;
    padding:14px 18px;
    border-radius:10px;
    margin:24px 0 14px 0;
">
<strong>🧠 String checkpoint</strong><br>
</div>

### Check Your Understanding 1 — String indexing and slicing

Given:

```python
text = "Data Analytics"
```

Store the substring `"Analytics"` in `your_ans` using **slicing**.

```python
text = "Data Analytics"
your_ans = ...
print(your_ans)
```

```python
grader.check("q1")
```

### Check Your Understanding 2 — String methods

The following text contains extra spaces and mixed capitalization:

```python
raw_text = "  PYTHON,DATA,AI  "
```

Use string methods to create:

```python
["python", "data", "ai"]
```

Store the result in `your_ans`.

```python
raw_text = "  PYTHON,DATA,AI  "
your_ans = ...
print(your_ans)
```

```python
grader.check("q2")
```

<div style="
    background:linear-gradient(135deg,#eff6ff,#eef2ff);
    border:1px solid #c7d2fe;
    padding:20px 24px;
    border-radius:14px;
    margin:30px 0 16px 0;
">
<h2 style="margin:0 0 8px 0;">📋 2. Lists</h2>
Lists are ordered and mutable. They are ideal when you need a collection whose contents may change.
</div>

In Python, lists are used to store items in a sequential order. These items can be of various types and are created using square brackets.

```python
fruits = ["Apple", "Banana", "Orange", "Kiwi"]
type(fruits)
```

```python
students = [["Kevin", 23, "484 Bourborn St."], ["Tyler", 34, "95 Esplanade Ave."]]
```

#### Indexing in Lists

Since lists are ordered, they maintain an index similar to strings. You can access items in a list using these indices.

```python
fruits[0]
```

```python
fruits[:3]
```

```python
students[-1]
```

### Check items in a list

You can check if an item is in a list, using `in` keyword

```python
"Kiwi" in fruits
```

```python
"Mango" in fruits
```

### Length of a list

To determine the number of items in a list, use the built-in `len()` function.

```python
len(students)
```

### Changing items in a list

Lists are mutable, meaning you can modify, add, or remove items. To change an item in a list, you can use indexing to assign a new value.

```python
fruits[0]= "Mango"
fruits
```

To add an item to the end of a list, use the `.append()` **method**. 

```python
fruits.append("Peach")
```

```python
fruits
```

### Adding items to a list

To insert an item at a particular index, use the `.insert()` **method**. 

```python
fruits.insert(0, "Apple")
fruits
```

To add all items from another iterable to a list, use the `.extend()` **method**.

```python
fruits.extend(["Pomegranate", "Pineapple"])
fruits
```

You may also use `+` to add/concatenate two lists together

```python
list1 = ["A", "B", "C"]
list2 = [1, 2, 3]
my_list = list1 + list2
my_list
```

### Remove items from a list

To remove an item by value, use the `.remove()` **method**.

```python
fruits.remove("Peach")
```

```python
fruits
```

To remove an item by index, use the `.pop()` **method**.

```python
fruits.pop(1)
fruits
```

```python
fruits.pop() # if no index is given, it will remove the last item from the list
fruits
```

### Copying list

When a list is copied to another list, both lists actually refer to the same items. Therefore, changes made to one list will be reflected in the other.

For eg: 

```python
list1 = [1, 2, 3]
list2 = list1 # here list1 is copied to list2

list2[0] = 100
list1 # although the first item of list2 is changed, however they both refer to the same items, it will reflect in list1
```

If you want to make a separate copy of the list, use `copy` method. 

```python
list1 = [1, 2, 3]
list2 = list1.copy()
list2[0] = 100
list1 # now changing list2 doesn't effect list1
```

```python
list2
```

For detailed list of other List methods, click [here](https://www.w3schools.com/python/python_ref_list.asp)

<div style="
    border-left:5px solid #3b82f6;
    background:#eff6ff;
    padding:14px 18px;
    border-radius:10px;
    margin:24px 0 14px 0;
">
<strong>🧠 List checkpoint</strong><br>
These questions check nested indexing, mutation, and the important difference between assignment and copying.
</div>

### Check Your Understanding 3 — Nested list indexing

Given the nested list:

```python
my_list = [56, [192, [395, 539, "abracadabra", 5]], [38, 59, 56]]
```

Access the string `"abracadabra"` and store it in `your_ans`.

```python
my_list = [56, [192, [395, 539, "abracadabra", 5]], [38, 59, 56]]
your_ans = ...
print(your_ans)
```

```python
grader.check("q3")
```

### Check Your Understanding 4 — Modify a list

Start with:

```python
scores = [78, 84, 91, 88]
```

Change the first score to `80`, append `95`, and store the resulting list in `your_ans`.

```python
scores = [78, 84, 91, 88]
scores[0] = ...
...
your_ans = ...
print(your_ans)
```

```python
grader.check("q4")
```

### Check Your Understanding 5 — Copying a list

Given:

```python
original = [1, 2, 3]
```

Create an **independent copy** named `copied`, then change the first item in `copied` to `100`.
The original list must remain unchanged.

```python
original = [1, 2, 3]
copied = ...
copied[0] = ...

print(original)
print(copied)
```

```python
grader.check("q5")
```

<div style="
    background:linear-gradient(135deg,#f0fdf4,#ecfdf5);
    border:1px solid #bbf7d0;
    padding:20px 24px;
    border-radius:14px;
    margin:30px 0 16px 0;
">
<h2 style="margin:0 0 8px 0;">📦 3. Tuples</h2>
Tuples are ordered like lists, but they are immutable. They are useful for fixed collections of values.
</div>

Like lists, tuples are used to store items in a sequential order. However, tuples are defined using parentheses and are immutable, meaning their contents cannot be changed after creation

```python
colors = ("Green", "Black", "Blue", "Red")
colors
```

```python
type(colors)
```

### Indexing in tuples

Similar to lists, you can use index in tuple to access items in a tuple. 

```python
colors[1]
```

```python
colors[:3]
```

```python
colors[-3]
```

Since tuples are immutable and their values are ordered, you can "unpack" each item of a tuple into individual variables as shown below:

```python
name = ("Dipak", "Singh")
first_name, second_name = name #Unpacking the two values of a tuple into two variables. 

print(f"First item: {first_name} ;  Second item: {second_name}")
```

For detailed list of other Tuple methods, click [here](https://www.w3schools.com/python/python_ref_tuple.asp)

<div style="
    border-left:5px solid #22c55e;
    background:#f0fdf4;
    padding:14px 18px;
    border-radius:10px;
    margin:24px 0 14px 0;
">
<strong>🧠 Tuple checkpoint</strong><br>
One unpacking exercise is enough to verify the main tuple skill introduced here.
</div>

### Check Your Understanding 6 — Tuple unpacking

Given:

```python
course = ("CSCI", "1462", "Data Analytics")
```

Unpack it into the variables `prefix`, `number`, and `title`.

```python
course = ("CSCI", "1462", "Data Analytics")
prefix, number, title = ...

print(prefix, number, title)
```

```python
grader.check("q6")
```

<div style="
    background:linear-gradient(135deg,#faf5ff,#fdf4ff);
    border:1px solid #e9d5ff;
    padding:20px 24px;
    border-radius:14px;
    margin:30px 0 16px 0;
">
<h2 style="margin:0 0 8px 0;">🗂️ 4. Dictionaries</h2>
Dictionaries organize data using key–value pairs, which makes them especially useful for named fields,
records, configuration data, and JSON-style data.
</div>

Dictionaries store data as **key–value pairs**. Dictionaries are mutable, and in Python 3.7+ they preserve insertion order. 

```python
my_info = {"name": "Steve", "age" : 30, "state" : "Texas"}
type(my_info)
```

```python
student = {"name": "Steve", "age" : 30, "courses" : ["CSCI1394", "CSCI924"]}
```

Dictionary keys must be **unique** and **hashable**. Strings, numbers, and tuples containing hashable values are common key types. 

```python
my_info  = {"name": "Steve", "age" : 30, "state" : "Texas", "name": "Kevin"}
```

```python
my_info
```

### Accessing items of a dictionary

You can access an item in a dictionary by referring to its key. 

```python
student["name"]
```

```python
student["courses"]
```

You can view all dictionary keys using the `.keys()` method.

```python
student.keys()
```

You can view all dictionary values using the `.values()` method.

```python
student.values()
```

You can view key–value pairs using the `.items()` method.

```python
student.items()
```

You can also access items in a nested dictionary like the following:

```python
nested_dict = {
    'person': {
        'name': 'John Doe',
        'age': 30,
        'address': {
            'street': '123 Main St',
            'city': 'Anytown',
            'zipcode': '12345'
        },
        'contacts': {
            'email': 'johndoe@example.com',
            'phone': '555-1234'
        }
    },
    'job': {
        'title': 'Software Engineer',
        'department': 'Engineering',
        'years_experience': 8
    }
}
```

```python
nested_dict["job"]["title"] # Accessing through multiple keys in dictionary
```

### Check if an item exists

To check if a key exists in a dictionary, you can use the `in` keyword

```python
"name" in student
```

### Changing items of a dictionary

To change a dictionary value, assign a new value to its key.

```python
student["name"] = "Tyler" # changing the name from Steve to Tyler
student
```

### Adding items to a dictionary

You can add an item to a dictionary like the following:

```python
student["country"] = "USA"
student
```

### Removing items from a dictionary

You can remove item from a dictionary by refering to the key in the `pop` function

```python
student.pop("country")
student
```

For detailed list of other Dictionary methods, click [here](https://www.w3schools.com/python/python_ref_dictionary.asp)

<div style="
    border-left:5px solid #a855f7;
    background:#faf5ff;
    padding:14px 18px;
    border-radius:10px;
    margin:28px 0 14px 0;
">
<strong>🧠 Dictionary checkpoint</strong><br>
These questions verify nested access and modification of key–value data.
</div>

### Check Your Understanding 7 — Nested dictionary access

Given:

```python
student = {
    "name": "Ana",
    "grades": {
        "Exam 1": 91,
        "Exam 2": 95
    }
}
```

Store the `Exam 2` score in `your_ans`.

```python
student = {
    "name": "Ana",
    "grades": {
        "Exam 1": 91,
        "Exam 2": 95
    }
}

your_ans = ...
print(your_ans)
```

```python
grader.check("q7")
```

### Check Your Understanding 8 — Modify a dictionary

Start with:

```python
inventory = {"apple": 3, "banana": 5}
```

Add `"orange"` with value `4`, remove `"apple"`, and store the final dictionary in `your_ans`.

```python
inventory = {"apple": 3, "banana": 5}
inventory["orange"] = ...
...
your_ans = ...
print(your_ans)
```

```python
grader.check("q8")
```

<div style="
    background: linear-gradient(135deg, #ecfdf5 0%, #eff6ff 100%);
    border:1px solid #c7e8df;
    padding:22px 26px;
    border-radius:14px;
    margin:30px 0 12px 0;
">
<h2 style="margin-top:0;">✅ Key Takeaways</h2>

### Strings
- Ordered and **immutable**
- Use indexing and slicing to access characters or substrings
- Common methods include `.upper()`, `.lower()`, `.count()`, `.strip()`, `.replace()`, and `.split()`

### Lists
- Ordered and **mutable**
- Use indexing and slicing to access values
- Common methods include `.append()`, `.insert()`, `.extend()`, `.remove()`, `.pop()`, and `.copy()`

### Tuples
- Ordered and **immutable**
- Support indexing, slicing, and unpacking
- Useful when a sequence should not be modified

### Dictionaries
- Store **key–value pairs**
- Mutable and accessed primarily by key
- Common methods include `.keys()`, `.values()`, `.items()`, and `.pop()`
- Can contain nested lists, tuples, or dictionaries

### Choosing a structure

| Need | Good choice |
|---|---|
| Store text | **String** |
| Editable ordered collection | **List** |
| Fixed ordered collection | **Tuple** |
| Named fields / key-based lookup | **Dictionary** |

</div>

Tests for q1 (points: None)
```
>>> your_ans == 'Analytics'
```


Tests for q2 (points: None)
```
>>> your_ans == ['python', 'data', 'ai']
```


Tests for q3 (points: None)
```
>>> your_ans == 'abracadabra'
```


Tests for q4 (points: None)
```
>>> your_ans == [80, 84, 91, 88, 95]
```


Tests for q5 (points: None)
```
>>> original == [1, 2, 3]
>>> copied == [100, 2, 3]
```


Tests for q6 (points: None)
```
>>> prefix == 'CSCI'
>>> number == '1462'
>>> title == 'Data Analytics'
```


Tests for q7 (points: None)
```
>>> your_ans == 95
```


Tests for q8 (points: None)
```
>>> your_ans == {'banana': 5, 'orange': 4}
```
