```python
# Initialize Otter
import otter
grader = otter.Notebook("KC2_3.ipynb")
```

Before working on the Knowledge Check, make sure you **run the Python code cells at the top of the notebook**. These cells initialize the autograder and are required for the tests to work correctly.

* If the first Python cell produces an error, try **restarting the kernel** by selecting **Kernel → Restart Kernel** from the JupyterLab menu.
* After restarting the kernel, run the first Python cell again, followed by the remaining setup cells.
* Make sure these setup cells run successfully before attempting the autograder tests.


# Knowledge Check 2 & 3

## Due: Sep 18 by 11:59 pm

## Instructions: 

1. After completing **KC2 and KC3**, download your completed `KC2_3.ipynb` file by selecting **File → Download** in JupyterLab.

2. Upload your completed `KC2_3.ipynb` file to the **KC2 Dropbox in D2L**.

3. This assignment is due by the **date and time listed in the course calendar**.

4. This is an **individual assignment**. As stated in the course syllabus, the use of AI tools is prohibited. You may not receive assistance from anyone other than the instructor, the TA, or the AARC. Any help received from the TA or AARC must be clearly documented.

5. If you use any resource other than the instructor or the course textbook, you must document it in a **new Markdown cell at the top of this notebook**. Include the source, what you learned from it, and specifically how you used that information in completing the assignment. Submissions that are copied, either partially or entirely, from another source will be subject to a significant penalty.

6. If you need help, visit the instructor during the posted office hours. If those times do not work with your schedule, email the instructor to arrange an alternative meeting time.

7. Some of the questions do not have autograder. You may have to test them by yourself. For autograded questions, passing the visible autograder tests does **not** guarantee that your answer is correct. The instructor may use additional hidden tests during grading. Full credit will be awarded only when your answers also pass these tests.

8. If you encounter any technical or assignment-related difficulties, please contact the instructor for assistance.

9. Reminder: If you successfully complete **at least 11 of the 13 Knowledge Checks (KCs)**, you will be eligible for **one additional opportunity to complete exam corrections**.


### Learning Objectives

By the end of this Knowledge Check, students should be able to:

* Conditional statements
* Read input from console in python

### Question 1: 

A shipping company uses the following function to calculate the cost (in dollars) of shipping based on the weight of the package (in pounds).

$$
c(w)=
\begin{cases}
3.5, & \text{if } 0 < w \le 1,\\
5.5, & \text{if } 1 < w \le 3,\\
8.5, & \text{if } 3 < w \le 10,\\
10.5, & \text{if } 10 < w \le 20.
\end{cases}
$$

Write a program that takes the weight of the package from user prompt and displays the shipping cost. If the weight is negative or zero, display a message `Invalid input`. If the weight is greater than 20, display a message `The package cannot be shipped.`

Test your program to check if it cover all cases.

Sample Output:

```python
Please enter the weight of the package and then hit Enter to run this cell:  5
The shipping cost is: 8.5

Please enter the weight of the package and then hit Enter to run this cell:  50
The package cannot be shipped

Please enter the weight of the package and then hit Enter to run this cell:  -5
Invalid input
```

```python
weight = int(input("Please enter the weight of the package and then hit Enter to run this cell: "))

if(weight < 0):
    ...
elif(...)
    shipping_cost = ...
    print(f"The shipping cost is: {shipping_cost}")
elif(...)
    ...
elif(...)
    ...
elif(...)
    ...
else:
    print("The package cannot be shipped")

```

### Question 2: 

Write a python program that prompts the user to enter the center x-y cordinates, width and height of two rectangles and determines whether the second rectangle is inside the first, overlaps, or do not overlap with the first, as shown in the figure below. 

<img src="tworectangle.png" alt="Example illustrating the different cases" width=500>

(a) A rectangle is inside another one;  (b) A rectangel overlaps another one

Your program should print as per the sample output runs for different cases:

```python
Enter the x cordinate, y cordinate, width and height of first rectangle separated by space 2.5 4 2.5 43
Enter the x cordinate, y cordinate, width and height of second rectangle separated by space 1.5 5 0.5 3
r2 is inside r1

Enter the x cordinate, y cordinate, width and height of first rectangle separated by space 1 2 3 5.5
Enter the x cordinate, y cordinate, width and height of second rectangle separated by space 3 4 4.5 5
r2 overlaps r1

Enter the x cordinate, y cordinate, width and height of first rectangle separated by space 1 2 3 3
Enter the x cordinate, y cordinate, width and height of second rectangle separated by space 40 45 3 2
r2 does not overlap r1
```


```python
rect1 = input("Enter the x cordinate, y cordinate, width and height of first rectangle separated by space")
rect1 = rect1.split()
x1 = float(rect1[0])
y1 = float(rect1[1])
w1 = float(rect1[2])
h1 = float(rect1[3])

rect2 = input("Enter the x cordinate, y cordinate, width and height of second rectangle separated by space")
rect2 = rect2.split()
x2 = float(rect2[0])
y2 = float(rect2[1])
w2 = float(rect2[2])
h2 = float(rect2[3])
```

```python
if(...):
    print("r2 is inside r1")
elif(...):
    print("r2 overlaps r1")
else:
    print("r2 does not overlap r1")
```

### Question 3: 

Write a program to compute the following summation:

$$
\frac{1}{3} + \frac{3}{5} + \frac{5}{7} + \frac{7}{9} + \frac{9}{11} + \frac{11}{13} + \cdots + \frac{95}{97} + \frac{97}{99}
$$

Store your answer in a variable `total` rounded upto two decimal points. 

```python
total = 0

for ... in range(...):
    ...

total = round(total, 2)
total
```

```python
grader.check("q3")
```

### Question 4: 

Assume that the letters A, E, I, O and U are vowels. Write a program that prompts the user to enter a string, and displays the number of vowels and consonants in the string.

Sample output would look like the following: 

```python
Enter a string: Programming is fun
The number of vowels is 5
The number of consonants is 11
```

```python
msg = input("Enter a string: ")
vowels = ['A', 'E', 'I', 'O', 'U']

count_of_vowels = 0
count_of_consonants = 0

...


print(f"The number of vowels is {count_of_vowels}")
print(f"The number of consonants is {count_of_consonants}")
```

### Question 5: 

Write a python program which prompts a user to enter a string, and finds the frequency of each word in the string. 
Store the final result in a variable, `your_ans` of dictionary (dict) type. 

Sample output:

```python
Enter a string: How much wood would a woodchuck chuck if a woodchuck could chuck wood
{'How': 1,
 'much': 1,
 'wood': 2,
 'would': 1,
 'a': 2,
 'woodchuck': 2,
 'chuck': 2,
 'if': 1,
 'could': 1
}

```

```python
msg = input("Enter a string: ")

your_ans = {}

...


your_ans
```

### Question 6: 

Write a program that prompts the user to enter a string and displays the string in reverse order. 

Sample output:

```python
Enter a string: ABCD
The reversed string is DBCA
```

```python
msg = input("Enter a string: ")

reversed_msg = ""

for letter in msg:
    reversed_msg = ...

print(f"The reversed string is {reversed_msg}")
```

### Question 7: 

Consider a nested list, `grid`. Write a python code to add all the elements of the list. 

Do NOT use any in-built functions such as `sum()` method. Store your answer in the variable `total`.

```python
grid = [
    [4, 2, 7, 2, 9],
    [6, 5, 3, 8, 1],
    [4, 9, 7, 6, 2],
    [8, 5, 3, 5, 7],
    [1, 6, 9, 8, 4]
]

total = 0

...


print(total)
```

### Question 8: 

Consider the following list, which consists of score of a student in a course during a semester. Write a python code to find the average score of the student in that semester. Store your answer in the variable `your_ans` rounded upto two decimal points. 

```python
scores_per_course = [('CSCI1462', 88), ('CSCI1302', 75), ('MATH1304', 80), ('ENGL1301', 90)]
```

```python
total = 0

...


your_ans = round(total/len(scores_per_course), 2)
your_ans
```

### Question 9: 

Write a program that takes two strings and find the largest common prefix of the two strings. Store your answer in the variable, `your_ans`. Consider the strings are case-sensitive. 

For eg: If first string is `Welcome to Python` and second string is `Welcome to Data Analytics`, then the largest common prefix is `Welcome to `. 

Similarly, If the first string is `Atlanta` and second string is `Macon`, the the largest common prefix is "" empty string. 


```python
msg1 = 'Welcome to Python'
msg2 = 'Welcome to Data Analytics'

your_ans = ""
```

```python
...



your_ans
```

### Question 10: 

The following dictionary contains the afternoon temperatures (in °F) recorded in several Texas cities:

```python
temperatures = {
    "Nacogdoches": 91,
    "Dallas": 95,
    "Houston": 93,
    "Austin": 97,
    "San Antonio": 96,
    "El Paso": 94,
    "Lubbock": 89,
    "Waco": 92
}
```

Write a Python program that uses a for loop to examine the temperatures and determine:

- The highest temperature recorded.
- The city where the highest temperature occurred.

Your program should produce output similar to:
```text
The highest temperature is 97 degrees Fahrenheit.
It was recorded in Austin.
```

Requirements:

- You must use a for loop to process the dictionary.
- Access each city's temperature using its dictionary key.
- Do not use the built-in max() function.
- Your program should work correctly even if the temperature values in the dictionary are changed.

*Hint*: Start by assuming the temperature of one city is the highest, and then compare it with the temperatures of the remaining cities.

```python
temperatures = {
    "Nacogdoches": 91,
    "Dallas": 95,
    "Houston": 93,
    "Austin": 97,
    "San Antonio": 96,
    "El Paso": 94,
    "Lubbock": 89,
    "Waco": 92
}
```

```python
max_temp = 91
max_city = "Nacogdoches"

...


print(f"The highest temperature is {max_temp} degrees Fahrenheit")
print(f"It was recorded in {max_city}")
```

Tests for q3 (points: None)
```
>>> total == 45.12
True
```
