```python
# Initialize Otter
import otter
grader = otter.Notebook("Numpy_Part1.ipynb")
```

# NumPy Arrays

So far, we have learned about different types of data structures in Python.

Some data structures are built directly into Python, such as:

- Strings
- Lists
- Tuples
- Dictionaries

Other data structures are provided by Python modules. For example, the `pandas` module provides:

- Series
- DataFrame

We have also learned that each data structure has methods associated with them that allow us to manipulate and work with the data.

For example:

```python
"hello".upper()
```
returns 
```python
'HELLO'
```
Similarly, 
```python
numbers = [5]
numbers.append(10)
numbers
```
returns
```python
[5, 10]
```

In this section, we are going to learn another important data structure:

# Numpy Arrays


## What is NumPy?

NumPy is a Python library that is mainly used for numerical and scientific computing.

Its main data structure is the **NumPy array**.

NumPy arrays are usually created using the `array()` function:

```python
numpy.array([...])
```

NumPy arrays are similar to Python lists because they allow us to store multiple values in an ordered sequence.

However, NumPy arrays have several important differences from Python lists.

We will explore these differences throughout this tutorial.

First make sure you import the `numpy` module to use numpy arrays. 


```python
import numpy as np
```

## Creating a NumPy Array

We can create a NumPy array using the `np.array()` function.

For example:


```python
numbers = np.array([10, 20, 30, 40, 50])

numbers
```

However, this is a NumPy array, not a Python list.

We can check its type using type().

```python
type(numbers)
```

**ndarray** stands for N-dimensional array.

This is the main data structure provided by NumPy.


## NumPy Arrays vs Python Lists

NumPy arrays and Python lists may look similar, but they behave differently.

Let's first create a Python list and a NumPy array containing the same values.


```python
my_list = [1, 2, 3, 4, 5]

my_array = np.array([1, 2, 3, 4, 5])

my_list, my_array
```

Both contain the same values, but they are different data structures.

Let's check their types.

```python
type(my_list), type(my_array)
```

The list is a Python `list`, while the NumPy object is a `numpy.ndarray`.

One of the biggest advantages of NumPy arrays is that they are designed specifically for numerical operations. Consider the following code which multiplies each number between 0 to 999999 using numpy methods and traditional for loop and python lists. 

```python
# you do not need to know how this code works. Keep reading the next cell
import time 

numbers_list = list(range(1_000_000))

start = time.time()
result = []
for number in numbers_list:
    result.append(number * 2)
end = time.time()
list_time = end - start


numbers_array = np.arange(1_000_000)
start = time.time()
result = numbers_array * 2
end = time.time()
numpy_time = end - start

print("List time:", list_time, "secs")
print("NumPy time:", numpy_time, "secs")

```

Your exact results will depend on your computer and what else is running on it.

However, NumPy will generally be much faster for large numerical operations because its operations are implemented using highly optimized compiled code.

This is especially important in data science because datasets can contain millions or even billions of values.

Instead of repeatedly writing Python `for` loops, we can often use NumPy's vectorized operations to process the data much more efficiently.

# Features of Numpy Array

## 1. Homogeneous Data Types

NumPy arrays have a single data type, known as their **dtype**.

For example, if we create an array containing integers:

```python
numbers = np.array([10, 20, 30, 40])

numbers
```

```python
numbers.dtype
```

The `dtype` tells us the type of data stored in the array. In the above example, numbers stores `int` type of values, and each value takes 64 bits in the memory.

We can also create an array containing real numbers.

```python
prices = np.array([10.5, 20.75, 30.25])

prices
```

```python
prices.dtype
```

Python lists can contain different types of data:

```python
my_list = [10, "hello", 3.5, True]
```

NumPy arrays, on the other hand, use a common data type for their elements.

```python
mixed = np.array([10, "hello", 3.5, True])

mixed
```

Notice that NumPy converted the integers, floats, booleans to string so that all elements could have the same data type.

This consistency allows NumPy to store and process data very efficiently.

## 2. NumPy Arrays Are Mutable

NumPy arrays are **mutable**, which means that we can change the values of existing elements.

For example:

```python
numbers = np.array([10, 20, 30, 40])

numbers
```

```python
numbers[0] = 100

numbers

```

The first element changed from `10` to `100`.

So NumPy arrays are mutable.

However, NumPy arrays generally have a **fixed size and shape** after they are created. This is different from Python lists, where we can easily append new elements.

## 3. Indexing and Slicing NumPy Arrays

Like Python lists, NumPy arrays use indexing.

Remember that Python indexing starts at `0`.

For example:

```python
numbers = np.array([10, 20, 30, 40, 50])

numbers[0]
```

```python
numbers[2]
```

```python
numbers[-1]
```

We can also use slicing, just like we do with Python lists.

```python
numbers[1:4]
```

```python
numbers[:3]
```

```python
numbers[2:]
```

## 4. Shape and Dimensions

One of the most important features of NumPy arrays is that they can have multiple dimensions.

A simple array such as:

```python
[10, 20, 30, 40]
```

is a one-dimensional array.

```python
numbers = np.array([10, 20, 30, 40])
numbers.ndim
```

The `ndim` attribute tells us the number of dimensions.

Our array has:

```text
1 dimension
```

We can also check its shape.

```python
numbers.shape
```

This tells us that the array contains 4 elements along its only dimension.

Now let's create a two-dimensional array.

```python
matrix = np.array([
    [1, 2, 3],
    [4, 5, 6]
])

matrix
```

```python
matrix.ndim
```

```python
matrix.shape
```

NumPy uses the shape of an array to understand how its data is organized. Here, (2, 3) means it's a 2-dimensional array, where there are two rows, and each row has 3 elements/columns. Each element in a NumPy array has a well-defined shape, which is crucial for performing vectorized operations. For instance, in a 2D array (matrix), all rows have the same number of columns, ensuring consistent dimensionality.

## 5. Accessing Elements in a 2D Array

For a two-dimensional array, we provide two indexes:

```python
array[row, column]
```

```python
matrix[0, 0]
```

```python
matrix[1, 1] # select the item at 2nd row and 2nd col of the matrix
```

```python
matrix[1, 1:] # select the second row, and columns starting from 2nd column to the very end
```

```python
matrix[1] # select the second row of the matrix
```

```python
matrix[:, 2] # select the third column of the matrix
```

```python
matrix[[0, 1], [2, 2]] # select the items at index (0, 2) and (1, 2) from the matrix 
```

# Why Do Data Scientists Use NumPy?

Many operations that we perform with NumPy can also be performed using Python lists and `for` loops.

So why do we use NumPy?

There are several reasons:

### 1. Speed

NumPy is highly optimized for numerical operations and can process large amounts of data very quickly.

### 2. Vectorized Operations

We can perform operations on entire arrays without explicitly writing `for` loops.

### 3. Memory Efficiency

NumPy arrays can store numerical data more efficiently than regular Python lists.

### 4. Multidimensional Data

NumPy makes it easy to work with one-dimensional, two-dimensional, and higher-dimensional data.

### 5. Foundation for Data Science

Many Python libraries used in data science and machine learning rely heavily on NumPy.

For example, pandas uses NumPy internally for many of its operations.

This is why learning NumPy is an important step toward learning data analysis.

# Knowledge Check 4 (Part 1) 

Based on what we have learned in the discussions above, complete the following questions and submit your work to the D2L KC4_Part1_Dropbox by 11:59 PM on September 22, 2026.

You do not need to answer every question correctly. As long as you make a genuine attempt and show your work, you can submit the assignment without any penalty.

Do NOT use `for` loops to answer any of the question. 

### Question 1: 

Create a numpy array to store the values 5, 10, 15 in a variable `arr`. 

*Hint*: See `np.array()` method above.

```python
arr = ...
arr
```

```python
grader.check("q1")
```

### Question 2: 

Create a numpy array to store the names of your favorite three fruits. Store your answer in a variable, `my_fruits`. 

```python
my_fruits = ...
my_fruits
```

```python
grader.check("q2")
```

<!-- BEGIN QUESTION -->

### Question 3: 

What is the cause of error in the following code snippet? 

Hint: Check the `shape` of each element of the numpy array.  Does each row of the numpy array has same number of columns?

```python
arr = np.array([[10, 20], [30]])
```

_Type your answer here, replacing this text._

<!-- END QUESTION -->

### Question 4

Create a 2-dimensional array which looks like the following: 

\begin{bmatrix}
10 & 20 \\
11 & 11 \\
100 & 200 \\
\end{bmatrix}

Store your answer in the variable, `your_ans`. What is the shape of the resultant array, `your_ans`? 

```python
your_ans = ...
your_ans
```

```python
grader.check("q4")
```

### Question 5

3-D arrays

An array that has 2-D arrays (matrices) as its elements is called a 3-D array. 

Store the following 3D array in a variable, `arr3d`.  What is the shape of the resultant array? 

<img src="pics/numpyTensors.jpg" style="width:300px" />

```python
arr3d = ...

arr3d
```

```python
arr3d.shape
```

```python
grader.check("q5")
```

### Question 6

You are given the following matrix, $A = \begin{bmatrix}
1 & 2 & 3 & 4 \\
10 & 20 & 30 & 40 \\
100 & 200 & 300 & 400 \\
-10 & -20 & -30 & -40 \\
-1 & -2 & -3 & -4 \\
\end{bmatrix}$


Your task is to write a python code to only access the elements of 3rd row (100, 200, 300, 400) from the matrix A. Your answer should be of shape (4, )

```python
A = np.array([
             [1, 2, 3, 4],
             [10, 20, 30, 40],
             [100, 200, 300, 400],
             [-10, -20, -30, -40],
             [-1, -2, -3, -4]
            ])

A
```

```python
A.shape
```

```python
your_ans = ...
your_ans
```

```python
grader.check("q6")
```

### Question 7

Given the array A (see Q. 6), write a python code to access first two rows of the matrix. Your answer should be $\begin{bmatrix}
1 & 2 & 3 & 4 \\
10 & 20 & 30 & 40 \\
\end{bmatrix}$ 
of shape (2, 4)

```python
your_ans = ...
your_ans
```

```python
your_ans.shape
```

```python
grader.check("q7")
```

### Question 8

Given the array A (see Q. 6), write a python code to access the last two rows of the matrix. Your answer should be $\begin{bmatrix}
-10 & -20 & -30 & -40 \\
-1 & -2 & -3 & -4 \\
\end{bmatrix}$ 
of shape (2, 4)

```python
your_ans = ...
your_ans
```

```python
your_ans.shape
```

```python
grader.check("q8")
```

### Question 9

Given the array A (see Q. 6), write a python code to access the first two columns of the matrix. Your answer should be $\begin{bmatrix}
1 & 2  \\
10 & 20  \\
100 & 200  \\
-10 & -20  \\
-1 & -2  \\
\end{bmatrix}$
of shape (5, 2)

```python
your_ans = ...
your_ans
```

```python
grader.check("q9")
```

### Question 10

Given the array A (see Q. 6), write a python code to access the first and the last rows of the matrix. Your answer should be 
$\begin{bmatrix}
1 & 2 & 3 & 4 \\
-1 & -2 & -3 & -4 \\
\end{bmatrix}$
of shape (2, 4)

```python
your_ans = ...
your_ans
```

```python
grader.check("q10")
```

### Question 11

Given the array A (see Q. 6), write a python code to access the first and the last columns of the matrix. Your answer should be 
$ \begin{bmatrix}
1  & 4 \\
10   & 40 \\
100  & 400 \\
-10 & -40 \\
-1  & -4 \\
\end{bmatrix}$
of shape (5, 2)

```python
your_ans = ...
your_ans
```

```python
your_ans.shape
```

```python
grader.check("q11")
```

### Question 12

Given the array A (see Q. 6), write a python code to access the element 200. 

```python
your_ans = ...
your_ans 
```

```python
grader.check("q12")
```

### Question 13

Given the square matrix $A = \begin{bmatrix}
1 & 2 & 3 \\
40 & 50 & 60  \\
-7 & -8 & -9  \\
\end{bmatrix}$. Write a python code to extract the diagonal elements of the matrix. Your answer should be a numpy array. 

```python
A = ...
your_ans = ...
your_ans
```

```python
grader.check("q13")
```

Tests for q1 (points: None)
```
>>> isinstance(arr, np.ndarray)
True
```
```
>>> list(arr) == [5, 10, 15]
True
```


Tests for q10 (points: None)
```
>>> your_ans.shape == (2, 4)
True
```
```
>>> your_ans[-1, -1].item() == -4
True
```


Tests for q11 (points: None)
```
>>> your_ans.shape == (5, 2)
True
```
```
>>> your_ans[-1, -1].item() == -4
True
```


Tests for q12 (points: None)
```
>>> your_ans.item() == 200
True
```


Tests for q13 (points: None)
```
>>> your_ans[-1].item() == -9
True
```
```
>>> isinstance(your_ans, np.ndarray)
True
```


Tests for q2 (points: None)
```
>>> isinstance(my_fruits, np.ndarray)
True
```
```
>>> len(my_fruits) == 3
True
```


Tests for q4 (points: None)
```
>>> your_ans.shape == (3, 2)
True
```


Tests for q5 (points: None)
```
>>> arr3d.shape == (3, 3, 3)
True
```


Tests for q6 (points: None)
```
>>> your_ans.shape == (4,)
True
```
```
>>> your_ans[-1].item() == 400
True
```


Tests for q7 (points: None)
```
>>> your_ans.shape == (2, 4)
True
```
```
>>> your_ans[-1, -1].item() == 40
True
```


Tests for q8 (points: None)
```
>>> your_ans.shape == (2, 4)
True
```
```
>>> your_ans[-1, -1].item() == -4
True
```


Tests for q9 (points: None)
```
>>> your_ans.shape == (5, 2)
True
```
```
>>> your_ans[-1, -1].item() == -2
True
```
