```python
# Initialize Otter
import otter
grader = otter.Notebook("Lab1.ipynb")
```

```python
import math
```

Before working on the Knowledge Check, make sure you **run the Python code cells at the top of the notebook**. These cells initialize the autograder and are required for the tests to work correctly.

* If the first Python cell produces an error, try **restarting the kernel** by selecting **Kernel → Restart Kernel** from the JupyterLab menu.
* After restarting the kernel, run the first Python cell again, followed by the remaining setup cells.
* Make sure these setup cells run successfully before attempting the autograder tests.

### Due Dates

**Core Skills:** Core Skill questions are due **during class**. If you are unable to complete all of them in class, you may submit the remaining questions with your Auxiliary Skill work. However, any Core Skill question completed after class can earn at most **50% of its available points**, even if the answer is correct.

**Auxiliary Skills:** Auxiliary Skill questions are due **September 4 by 11:59 p.m.**


## Laboratory Grading Policy

Each lab includes two types of tasks:

* **Core Tasks** – demonstrate your understanding of the essential skills and must be completed during the scheduled lab period.
* **Auxiliary Tasks** – build on the core skills and may usually be completed after class.

### Lab Marks

| Mark  |                                       Core Tasks |                                  Auxiliary Tasks |
| ----- | -----------------------------------------------: | -----------------------------------------------: |
| **E** |                                     At least 90% |                                     At least 90% |
| **M** |                                     At least 80% |                                     At least 70% |
| **N** |                                     At least 70% |                                     At least 60% |
| **X** | Below the N requirement or insufficient evidence | Below the N requirement or insufficient evidence |

### Overall Laboratory Grade

There are **six laboratory assignments**.

| Course Grade | Required Lab Performance                                            |
| ------------ | ------------------------------------------------------------------- |
| **A**        | **E** on at least 5 labs; the remaining lab must be **N or higher** |
| **B**        | **M or higher** on at least 5 labs; the remaining lab may be **X**  |
| **C**        | **N or higher** on at least 5 labs                                  |
| **D**        | **N or higher** on at least 4 labs                                  |
| **F**        | Does not meet the requirements for a D                              |

### Lab Completion Rules

**Approved Absences and Accommodations:** The 50% Core Task limit does not apply to approved institutional absences, instructor-approved excused absences, or applicable accommodations. A make-up lab or practical demonstration will be provided when appropriate.

**One Lab Reassessment:** You may reassess **one lab marked N or X**. The reassessment must be completed within **7 calendar days after feedback is provided** and may include a corrected submission, a practical demonstration, or both. The revised mark may be **M or E** if you independently demonstrate the required level of proficiency.


## Instructions: 

1. After completing **Lab1**, download your completed `Lab1.ipynb` file by selecting **File → Download** in JupyterLab.

2. Upload your completed `Lab1.ipynb` file to the **Lab1 Dropbox in D2L**.

3. This assignment is due by the **date and time listed in the course calendar**.

4. This is an **individual assignment**. As stated in the course syllabus, the use of AI tools is prohibited. You may not receive assistance from anyone other than the instructor, the TA, or the AARC. Any help received from the TA or AARC must be clearly documented.

5. If you use any resource other than the instructor or the course textbook/notes, you must document it in a **new Markdown cell in Question 20 of this notebook**. Include the source, what you learned from it, and specifically how you used that information in completing the assignment. Submissions that are copied, either partially or entirely, from another source will be subject to a significant penalty.

6. If you need help, visit the instructor during the posted office hours. If those times do not work with your schedule, email the instructor to arrange an alternative meeting time.

7. Passing the visible autograder tests does **not** guarantee that your answer is correct. The instructor may use additional hidden tests during grading. Full credit will be awarded only when your answers also pass these tests.

8. If you encounter any technical or assignment-related difficulties, please contact the instructor for assistance.


# Core Skills (Due in-class)

### Question 1: 
Consider the following code:

```python
temperatures = [72, 75, 71, 68, 74]
```

Write Python code to determine the total number of elements in the list and store the result in the variable `your_ans`.

Your solution should work for any list, not just `temperatures`.

```python
temperatures = [72, 75, 71, 68, 74]
your_ans = ...
```

```python
grader.check("q1")
```

### Question 2:

Write Python code to count the number of times the substring `th` appears in the following `text`. Store the result in `your_ans`.

```python
text = "The weather today is better than the weather yesterday."
```


```python
text = "The weather today is better than the weather yesterday."
your_ans = ...
your_ans
```

```python
grader.check("q2")
```

<!-- BEGIN QUESTION -->

### Question 3:

Given:

```python
scores = [82, 91, 76, 88]
```

What is the value of scores[-1]?

_Type your answer here, replacing this text._

<!-- END QUESTION -->

### Question 4: 

Consider the following empty list, `fruits`. 

```python
fruits = []
```

Add the elements `Apple`, `Banana`, `Mango` to the list. 

```python
fruits = []
fruits = ...
fruits
```

```python
grader.check("q4")
```

### Question 5: 

Add `"Kiwi"` as the **second element** in the `fruits` list.

**Important:** If your solution does not pass the grader and you modify the `fruits` list while trying different answers, re-run the code from **Question 4** before trying again. This will reset `fruits` to:

```python
["Apple", "Banana", "Mango"]
```

Otherwise, previous attempts may leave extra or duplicate items in the list, which can cause the grader to fail even after your code is correct.


```python
...
fruits
```

```python
grader.check("q5")
```

### Question 6: 

Remove the item `Mango` from the list. 

```python
...
fruits
```

```python
grader.check("q6")
```

### Question 7: 

Given a person's full name containing a first name, middle initial, and last name, write Python code to extract each part and store them in separate variables.

Your solution should work for any name that follows this format.

```python
person_name = "John K Doe"
first_name = ...
mi = ...
last_name = ...

print(f"First Name: {first_name}\nMI: {mi}\nLast Name: {last_name}")
```

```python
grader.check("q7")
```

### Question 8:

Consider the following empty dictionary, `actor`.

Add the following key-value pairs to the dictionary:

* `name` with the value `Tom Cruise`
* `last5` with the value `[150, 100, 100, 70, 50]`


```python
actor = {} 
actor['name'] = ...
actor['last5'] = ...
actor
```

```python
grader.check("q8")
```

<!-- BEGIN QUESTION -->

### Question 9: 

What is the output of the following python code?

```python
record = ("Alice", 21, "Data Analytics", 3.8)
x, y, z = record[1:4]
print(y)
```

_Type your answer here, replacing this text._

<!-- END QUESTION -->

### Question 10: 

Which of the following data structures are immutables?

A. Strings <br>
B. Lists <br>
C. Tuples <br>
D. Dictionaries <br>

Write your options(separated by comma and no space) in the `your_ans` variables. Do not remove the double quotes, or else, your test will fail. 

For eg: If the answers are option A, C, and E then your answer variable will look like `your_ans = "A,C,E"`

```python
your_ans = ...
```

```python
grader.check("q10")
```

# Auxiliary Skills (Due Sep 4 by 11:59 pm)

### Question 11: 

Write Python expression in the next cell that evaluates the following expression and save the output in a variable named _result_:
$$2 * (5\frac{5}{2}) + 2^{0.5 * 22} - \frac{7}{45} + 1$$

Please note by $a\frac{b}{c}$,  I mean $a + \frac{b}{c}$ and not $a * \frac{b}{c}$

```python
result = ...
print(result)
```

```python
grader.check("q11")
```

### Question 12: 

In class, we learned how to show output on the screen using the `print` function. For example, to display the message **I love Data Analytics**, you would use the following `print` statement:

```python
print("I love Data Analytics")
```

In the same way, you can use the `print` statement to show the value of a variable, as demonstrated below:

```python
my_var = 10
print(my_var)
```

In this question, we'll explore how to display both strings and variables together on the screen. Often, rather than just printing the value of a variable, we might want to format that value within a string. For example, take a look at the following cell:

```python
value_of_pi = round(math.pi, 2) # stores the value of pi rounded upto two decimal points
print(value_of_pi) 
```

If we want to format the value of pi with the message **The value of pi is 3.14**, we can use Python's [formatted string literal](https://docs.python.org/3/tutorial/inputoutput.html#tut-f-strings) for this purpose.  

Formatted string literals (also called f-strings for short) let you include the value of Python expressions inside a string by prefixing the string with f or F and writing expressions as `{expression}`

For eg: 

```python
print(f"The value of pi is {value_of_pi}")
```

Notice that the value_of_pi variable is enclosed in curly braces within the string so that it can be replaced by the value it holds when using the print statement. 

#### Task: 

Mark is employed at a Data Analytics firm with a yearly salary of $80,000 and receives a 5% bonus for 2023. To determine Mark's total gross income for the end of the year, write a Python code to compute the total amount, including the bonus. Also, display the gross income in the specified format:

`Mark's total gross income is <total_gross_income> dollars in 2023.`

Note: Ensure that you replace `<total_gross_income>` with the variable representing the calculated income. Avoid inserting the amount directly; instead, use the appropriate variable in the print statement.

Your answer should look like the following: 

`Mark's total gross income is 84000.0 dollars in 2023.` 

Note, your code should work for any salary of Mark. 



```python
mark_salary = ...
total_income = ...

print(...)
```

```python
grader.check("q12")
```

<!-- BEGIN QUESTION -->

### Question 13: 

Kevin, who is new to Python programming, is attempting to write a python code to display a total number of seconds into minutes and remaining seconds, and print a message in the format:

`<seconds> seconds is equivalent to <minutes> minutes and <seconds> remaining seconds.`

For eg: `125 seconds is equivalent to 2 minutes and 5 seconds`.


However, Kevin's function contains **syntax, runtime or logical errors** and isn't producing the correct output. 

List all the **syntax and runtime errors** that Kevin has made in the following code in each line. 

Use the following format for your answer:

Line 1: _list all errors_ <br>
Line 2: _list all errors_ <br>
Line 3: _list all errors_ <br>

If a line contains no error, write:

Line X: No error

```python
secs = 125
mins = seconds / mins
    remaining_secs = seconds // 0
    print(f"125 seconds is equivalent to {mins} mins and {secs} rems")
```

_Type your answer here, replacing this text._

<!-- END QUESTION -->

### Question 14: 

Fix the code in Question 13, so that Kevin achieves his desired output. 

Your code should work for any value of seconds.

```python
seconds = 125
mins = ...
remaining_secs = ...
...
```

```python
grader.check("q14")
```

### Question 15: 

Write a Python code to add 600 after 500 in the list provided below. Use the `insert` method to add 600, rather than manually inserting it.

```python
my_list = [1, 2, [30, 40, [500, 700], 40, 30], 2, 1]
```

Your final list, `my_list`, should appear as follows:

`[1, 2, [30, 40, [500, 600, 700], 40, 30], 2, 1]`

```python
my_list = [1, 2, [30, 40, [500, 700], 40, 30], 2, 1]
...
my_list
```

```python
grader.check("q15")
```

### Question 16: 

In the 2010s, Cambridge Analytica, a data analytics firm, collected data from millions of Facebook users through an application **without their consent** and aimed to use this information for electoral manipulation. Data from social media can reveal extensive details about an individual's lifestyle, beliefs, and personal information, potentially influencing their political views.

As data analysts, it is important to be cautious of such unethical uses of data.

In this task, you will explore how to extract sensitive information from social media data, though in a ethical manner, since the tweet used in this question is publicly available.

You have a sample of tweets from Donald Trump during his presidency (2017-2021). Your goal is to identify and extract the Twitter user screen name mentioned in his tweets.

Your code should return the string, `henrymcmaster`.

Hint: Use the `user_mentions` key to access the required information.

```python
trump_tweet ={"created_at":["Tue Jun 26 01:35:03 +0000 2018"],"id":[1.01142255594771e+18],"id_str":["1011422555947712513"],"full_text":["Thank you South Carolina. Now let’s get out tomorrow and VOTE for @HenryMcMaster! https://t.co/5xlz0wfMfu"],"display_text_range":[[0],[105]],"entities":{"hashtags":[],"symbols":[],"user_mentions":[{"screen_name":["henrymcmaster"],"name":["Henry McMaster"],"id":[16335288],"id_str":["16335288"],"indices":[[66],[80]]}],"urls":[{"url":["https://t.co/5xlz0wfMfu"],"expanded_url":["https://www.pscp.tv/w/bgGOtTFvTlFsTFJub1dwUXd8MXlvSk1WZHJWQm54Uf-J8fPu1RO4E84ax-LuK1bAbiCpnzBBZmdPfI9FAhGV?t=11s"],"display_url":["pscp.tv/w/bgGOtTFvTlFs…"],"indices":[[82],[105]]}]},"source":[r"<a href=\"http://twitter.com/download/iphone\" rel=\"nofollow\">Twitter for iPhone<\/a>"], "lang":["en"]}
trump_tweet
```

```python
# Do not simply type the name in the blanks. Write a python code to extract that information. 
trump_mentions = ...
trump_mentions
```

```python
grader.check("q16")
```

### Question 17: 

Consider:
```python
products = {
    "Laptop": 950,
    "Tablet": 600,
    "Monitor": 300
}
```

Without manually typing 950, write Python code to determine the **highest price** among the products. Store your result in the variable `your_ans`.

```python
products = {
    "Laptop": 950,
    "Tablet": 600,
    "Monitor": 300
}
your_ans = ...
your_ans
```

```python
grader.check("q17")
```

### Question 18: 

A data analyst collected the following tweet for sentiment analysis:

```python
tweet = "   WOW!!! 🤩 Python makes Data Analytics SO much fun!!! #DataAnalytics   "
```

Clean the tweet so that the final result becomes:

`wow python makes data analytics so much fun dataanalytics`

Store your final result in the variable `clean_tweet`

_Note_: Your program should work for any tweet that requires the same cleaning steps, including removing extra spaces, exclamation marks (!), laughing emojis, and hashtags (#).

```python
tweet = "   WOW!!!🤩 Python makes Data Analytics SO much fun!!! #DataAnalytics   "
clean_tweet = ...

print(clean_tweet)
```

```python
grader.check("q18")
```

### Question 19: 

Newton's law of universal gravitation states that every particle attracts every other particle in the universe with a force (in Newton) that is given as follows:
$$F = \frac{G * m_1 * m_2}{d^2}$$
Where G is the gravitational constant, $m_1$ and $m_2$ are the mass of the two objects, and d is the distance between them. 

Write a python code which calculates the gravitational force with which the moon attracts that objects (of mass 160 lbs) towards it according to the formula above. 

Given $G = 6.674 * 10^{-11} Nm^2{kg}^2$, 
mass of moon $ = 6 * 10^{24} kg$ and 
distance between moon and earth's surface is $3.84 * 10^8 m$

Note: 1 lbs is equal to 0.453 kg

```python
G = ...
mass_of_moon = ...
mass_of_object = ...
distance = ...
F = ...
F
```

```python
grader.check("q19")
```

<!-- BEGIN QUESTION -->

### Question 20: 

List all external sources (excluding the lecture notes and textbook), you used to answer the quiz questions, detailing the source and describing exactly what was learned and how that information was used in your assignment.

<!-- END QUESTION -->



Tests for q1 (points: None)
```
>>> your_ans > 0
True
```


Tests for q10 (points: None)
```
>>> your_ans.split(',')[0] in 'A,B,C,D'
True
```
```
>>> your_ans.split(',')[1] in 'A,B,C,D'
True
```


Tests for q11 (points: None)
```
>>> math.isclose(result, 2063.844, abs_tol=0.1)
True
```


Tests for q12 (points: None)
```
>>> int(total_income) == 84000
True
```


Tests for q14 (points: None)
```
>>> mins == 2
True
```
```
>>> remaining_secs == 5
True
```


Tests for q15 (points: None)
```
>>> my_list[2][2][1] == 600
True
```


Tests for q16 (points: None)
```
>>> trump_mentions == 'henrymcmaster'
True
```


Tests for q17 (points: None)
```
>>> your_ans == 950
True
```


Tests for q18 (points: None)
```
>>> clean_tweet == 'wow python makes data analytics so much fun dataanalytics'
True
```


Tests for q19 (points: None)
```
>>> math.isclose(F, 0.2, abs_tol=0.1)
True
```


Tests for q2 (points: None)
```
>>> your_ans > 0
True
```


Tests for q4 (points: None)
```
>>> len(fruits) == 3
True
```
```
>>> fruits[0] == 'Apple'
True
```


Tests for q5 (points: None)
```
>>> fruits[1] == 'Kiwi'
True
```
```
>>> len(fruits) == 4
True
```


Tests for q6 (points: None)
```
>>> len(fruits) == 3
True
```
```
>>> fruits[-1] == 'Banana'
True
```


Tests for q7 (points: None)
```
>>> first_name == 'John'
True
```
```
>>> mi == 'K'
True
```
```
>>> last_name == 'Doe'
True
```


Tests for q8 (points: None)
```
>>> actor['name'] == 'Tom Cruise'
True
```
```
>>> len(actor['last5']) == 5
True
```
