```python
# Initialize Otter
import otter
grader = otter.Notebook("Lab2.ipynb")
```

```python
import math
import pandas as pd
```

Before working on the Lab, make sure you **run the Python code cells at the top of the notebook**. These cells initialize the autograder and are required for the tests to work correctly.

* If the first Python cell produces an error, try **restarting the kernel** by selecting **Kernel → Restart Kernel** from the JupyterLab menu.
* After restarting the kernel, run the first Python cell again, followed by the remaining setup cells.
* Make sure these setup cells run successfully before attempting the autograder tests.

### Due Dates

**Core Skills:** Core Skill questions are due **during class**. If you are unable to complete all of them in class, you may submit the remaining questions with your Auxiliary Skill work. However, any Core Skill question completed after class can earn at most **50% of its available points**, even if the answer is correct.

**Auxiliary Skills:** Auxiliary Skill questions are due **September 18 by 11:59 p.m.**

## Laboratory Grading Policy

Each lab includes two types of tasks:

* **Core Tasks** – demonstrate your understanding of the essential skills and must be completed during the scheduled lab period.
* **Auxiliary Tasks** – build on the core skills and may usually be completed after class.

Each question carries **10 points**.

### Lab Marks

| Mark  |                                       Core Tasks |                                  Auxiliary Tasks |
| ----- | -----------------------------------------------: | -----------------------------------------------: |
| **E** |                                     At least 90% |                                     At least 90% |
| **M** |                                     At least 80% |                                     At least 70% |
| **N** |                                     At least 70% |                                     At least 60% |
| **X** | Below the N requirement or insufficient evidence | Below the N requirement or insufficient evidence |

### Overall Laboratory Grade

There are **six laboratory assignments**.

| Course Grade | Required Lab Performance                                            |
| ------------ | ------------------------------------------------------------------- |
| **A**        | **E** on at least 5 labs; the remaining lab must be **N or higher** |
| **B**        | **M or higher** on at least 5 labs; the remaining lab may be **X**  |
| **C**        | **N or higher** on at least 5 labs                                  |
| **D**        | **N or higher** on at least 4 labs                                  |
| **F**        | Does not meet the requirements for a D                              |

### Lab Completion Rules

**Approved Absences and Accommodations:** The 50% Core Task limit does not apply to approved institutional absences, instructor-approved excused absences, or applicable accommodations. A make-up lab or practical demonstration will be provided when appropriate.

**One Lab Reassessment:** You may reassess **one lab marked N or X**. The reassessment must be completed within **7 calendar days after feedback is provided** and may include a corrected submission, a practical demonstration, or both. The revised mark may be **M or E** if you independently demonstrate the required level of proficiency.


## Instructions: 

1. After completing **Lab2**, download your completed `Lab1.ipynb` file by selecting **File → Download** in JupyterLab.

2. Upload your completed `Lab2.ipynb` file to the **Lab2 Dropbox in D2L**.

3. This assignment is due by the **date and time listed in the course calendar**.

4. This is an **individual assignment**. As stated in the course syllabus, the use of AI tools is prohibited. You may not receive assistance from anyone other than the instructor, the TA, or the AARC. Any help received from the TA or AARC must be clearly documented.

5. If you use any resource other than the instructor or the course textbook/notes, you must document it in a **new Markdown cell in Question 15 of this notebook**. Include the source, what you learned from it, and specifically how you used that information in completing the assignment. Submissions that are copied, either partially or entirely, from another source will be subject to a significant penalty.

6. If you need help, visit the instructor during the posted office hours. If those times do not work with your schedule, email the instructor to arrange an alternative meeting time.

7. Passing the visible autograder tests does **not** guarantee that your answer is correct. The instructor may use additional hidden tests during grading. Full credit will be awarded only when your answers also pass these tests.

8. If you encounter any technical or assignment-related difficulties, please contact the instructor for assistance.


# Core Skills (Due in-class)

### Question 1:

Given a list of strings, count how many times each string appears in the list. Store the results in a dictionary named `your_ans`, where each string is a key and its frequency is the corresponding value.

For eg: 
```python
animals = ['dog', 'cat', 'elephant', 'cat', 'dog', 'dog', 'dog', 'fox', 'elephant']
```

Your answer should be: 

```
{ 'dog': 4, 'cat': 2, 'elephant': 2, 'fox': 1 }
```

```python
your_ans = {}

animals = ['dog', 'cat', 'elephant', 'cat', 'dog', 'dog', 'dog', 'fox', 'elephant']

for animal in animals:
    if(animal not in your_ans):
        your_ans[animal] = 1
    else:
        your_ans[animal] += 1


your_ans
```

Output:
```
{'dog': 4, 'cat': 2, 'elephant': 2, 'fox': 1}
```

```python
grader.check("q1")
```

Output:
```
q1 results: All test cases passed!
```

### Question 2:

Suppose a passing score for a course is 60 out of 100. The variable `scores` contains the exam scores of all students in the course.

Write a Python program that uses only **one for loop** to determine:
- the highest score,
- the lowest score, and
- the total number of students who passed the exam.

For eg:

```python
scores = [55, 59, 71, 62, 40, 90, 12, 85, 58, 22, 91]
```

Your output should be: 

```python
Highest: 91 Lowest: 12
Total Passed: 5
```

Restriction: **Do not use** python in-built functions `max()`, `min()`, `sum()`

```python
scores = [55, 59, 71, 62, 40, 90, 12, 85, 58, 22, 91]

highest  = 0
lowest = 100
total_passed = 0

for score in scores:
    if(score > highest):
        highest = score

    if(score < lowest):
        lowest = score

    if(score >= 60):
        total_passed += 1

print(f"Highest: {highest} Lowest: {lowest}")
print(f"Total Passed: {total_passed}")

```

Output:
```
Highest: 91 Lowest: 12
Total Passed: 5
```

```python
grader.check("q2")
```

Output:
```
q2 results: All test cases passed!
```

### Question 3:

Write a Python program to calculate the sum of the main diagonal elements of a square matrix. Store your answer in the variable `your_ans`.

A square matrix has the same number of rows and columns.

For example:
```python
matrix = [
    [4, 2, 7],
    [1, 5, 3],
    [8, 6, 9]
]
```
The main diagonal elements are:
```text
4, 5, 9
```
Your answer should be `18`

_Hint_: Find out what the row and column indices are for all the diagonal elements. Do you find any pattern? Are both of them same?

```python
matrix = [
            [4, 2, 7],
            [1, 5, 3],
            [8, 6, 9]
        ]
```

```python
your_ans = 0

for i in range(len(matrix)):
    your_ans += matrix[i][i]

your_ans
```

Output:
```
18
```

```python
grader.check("q3")
```

Output:
```
q3 results: All test cases passed!
```

### Question 4:

Create a Pandas Series from the following student scores:

```python
scores = [78, 92, 65, 88, 55]
students = ["Alice", "Bob", "Charlie", "David", "Emma"]
```

Write Python code to create a Series named `student_scores` using the student names as the index.

```python
scores = [78, 92, 65, 88, 55]
students = ["Alice", "Bob", "Charlie", "David", "Emma"]

student_scores = pd.Series(scores, index = students)
student_scores
```

Output:
```
Alice      78
Bob        92
Charlie    65
David      88
Emma       55
dtype: int64
```

```python
grader.check("q4")
```

Output:
```
q4 results: All test cases passed!
```

### Question 5:

Use `pd.DataFrame` method to create a dataframe, `fruits`, which looks like the following:

<img src="pics/fruits.png" width=200\>


Hint: Use `index` paramter of the `pd.DataFrame` method to add row named labels, similar to pandas series.


```python
fruits_name = ['Orange', 'Apple', 'Mango', 'Apricot', 'Strawberry']
quantity = [10, 12, 5, 6, 20]
price = [1.25, 4.25, 1.00, 3.22, 1.99]

season = ['Winter', 'Fall', 'Summer', 'Fall', 'Summer']

fruits = pd.DataFrame({'Fruit Name': fruits_name, 'Quantity': quantity, 'Price per lbs': price}, index = season)

fruits
```

Output:
```
Fruit Name  Quantity  Price per lbs
Winter      Orange        10           1.25
Fall         Apple        12           4.25
Summer       Mango         5           1.00
Fall       Apricot         6           3.22
Summer  Strawberry        20           1.99
```

```python
grader.check("q5")
```

Output:
```
q5 results: All test cases passed!
```

### Question 6:

Write a Python program (using **`loc`** property) to display the `Fruit Name` and `Price per lbs` for fruits that are typically grown in the summer season. Store your answer in a variable, `summer_fruits` 

Your solution would look like the following:

<img src="pics/summer_fruits.png" width=200/>

```python
summer_fruits = fruits.loc['Summer' , ['Fruit Name', 'Price per lbs']]
summer_fruits
```

Output:
```
Fruit Name  Price per lbs
Summer       Mango           1.00
Summer  Strawberry           1.99
```

```python
grader.check("q6")
```

Output:
```
q6 results: All test cases passed!
```

### Question 7:

Solve the question in 1.1(b) using **`iloc`** property. 

```python
summer_fruits = fruits.iloc[ [2, 4], [0, 2]]
summer_fruits
```

Output:
```
Fruit Name  Price per lbs
Summer       Mango           1.00
Summer  Strawberry           1.99
```

```python
grader.check("q7")
```

Output:
```
q7 results: All test cases passed!
```

### Question 8:

Given a list of person's name and age as shown below, find the name of the person with the youngest age. Assume no two people have same name or age.

For eg:

```python
persons = [('Kevin', 23), ('Sarah', 24), ('Justin', 21), ('Peter', 24), ('Chris', 32)]
```

The youngest person is `Chris`

```python
persons = [('Kevin', 23), ('Sarah', 24), ('Justin', 21), ('Peter', 24), ('Chris', 32)]
```

```python
youngest = persons[0][0]
min_age = persons[0][1]

for name, age in persons:
    if(age < min_age):
        min_age = age
        youngest = name

print(f"The youngest person is: {youngest}")
```

Output:
```
The youngest person is: Justin
```

```python
grader.check("q8")
```

Output:
```
q8 results: All test cases passed!
```

# Auxiliary Skills (Due by Sep 18 by 11:59 pm)

### Question 9:

Write a Python program to calculate the sum of the elements row-wise of a matrix of `nrows` rows and `ncols` columns. Store your answer in the variable `your_ans`.


For example:
```python
matrix = [
    [4, 2, 7, 1, 1],
    [1, 5, 3, 2, 5],
    [8, 6, 9, 1, 0]
]
```
The column-wise sum of elements are:
```text
4+2+7+1+1 = 15
1+5+3+2+5 = 16
8+6+9+1+0 = 24
```
Your answer should be `[15, 16, 24]`

Restriction: Do not use python in-built `sum` function

_Hint_: Use nested loops

```python
matrix = [
    [4, 2, 7, 1, 1],
    [1, 5, 3, 2, 5],
    [8, 6, 9, 1, 0]
]
```

```python
nrows = len(matrix) # total number of rows
ncols = len(matrix[0]) # total number of cols

your_ans = []

for r in range(nrows):
    total = 0
    for c in range(ncols):
        total += matrix[r][c]

    your_ans.append(total)


your_ans
```

Output:
```
[13, 9, 23]
```

```python
grader.check("q9")
```

Output:
```
q9 results:
    q9 - 1 result:
        ❌ Test case failed
        Trying:
            your_ans == [15, 16, 24]
        Expecting:
            True
        **********************************************************************
        Line 1, in q9 0
        Failed example:
            your_ans == [15, 16, 24]
        Expected:
            True
        Got:
            False
```

### Question 10:

Write a Python program to calculate the sum of the elements column-wise of a matrix of `nrows` rows and `ncols` columns. Store your answer in the variable `your_ans`.


For example:
```python
matrix = [
    [4, 2, 7, 1, 1],
    [1, 5, 3, 2, 5],
    [8, 6, 9, 1, 0]
]
```
The column-wise sum of elements are:
```text
4+1+8 = 13
2+5+6 = 13
7+3+9 = 19
1+2+1 = 4
1+5+0 = 6
```
Your answer should be `[13, 13, 19, 4, 6]`

Restriction: Do not use python in-built `sum` function

_Hint_: Use nested loop

```python
matrix = [
    [4, 2, 7, 1, 1],
    [1, 5, 3, 2, 5],
    [8, 6, 9, 1, 0]
]
```

```python
nrows = len(matrix) # total number of rows
ncols = len(matrix[0]) # total number of cols

your_ans = []

...


your_ans
```

```python
grader.check("q10")
```

### Question 11:

A table named `instructor` contains one row for each SFA faculty member and these columns:
- **Dept**: a string, the department of the faculty member
- **Name**: a string, the first name fo the faculty member
- **Designation**: a string, one of: Lecturer, Assitant Professor, Associate Professor, or Professor
- **Salary**: an int, last year's salary paid by the university

Write a python code to display the **last three rows** of the dataframe. 

_Hint_: Read pandas `tail()` method

```python
instructor = pd.read_csv('data/instructor.csv')
```

```python
instructor.tail(3)
```

Output:
```
Dept     Name         Designation  Salary
497    Nursing    Rylee            Lecturer   96354
498   Aviation  Jasmine  AssociateProfessor  113629
499  Sociology   Aurora  AssociateProfessor   66527
```

### Question 12:

What is the highest paying salary for an instructor across the university? Store your answer in the variable, `your_ans`.

_Hint_: You may use the python in-built `max()` function

```python
your_ans = max(instructor['Salary'])
```

```python
grader.check("q12")
```

Output:
```
q12 results: All test cases passed!
```

### Question 13:

What is the name of the instructor who earns the highest salary at the university? Store the instructor's name in the variable `your_ans`.

_Hint_: First, use a for loop to identify the row index corresponding to the highest salary. Then, use the `.iloc` property to retrieve the instructor's name from that row.

```python
your_ans = ""
max_salary = instructor.iloc[0, 3]
row_ix = 0

for salary in instructor['Salary']:

    if(salary > max_salary):
        max_salary = salary
        your_ans = instructor.iloc[row_ix, 1]
        
    row_ix += 1

your_ans
```

Output:
```
'Bella'
```

```python
grader.check("q13")
```

Output:
```
q13 results: All test cases passed!
```

### Question 14:

Write a python code to compute the following sum:

$$
\frac{1}{\sqrt{1} + \sqrt{2}} + \frac{1}{\sqrt{2} + \sqrt{3}} + \frac{1}{\sqrt{3} + \sqrt{4}} + \cdots + \frac{1}{\sqrt{624} + \sqrt{625}}
$$

Store your answer in the variable `your_ans` rounded upto two decimal points.

```python
your_ans = 0

for i in range(1, 625):
    your_ans += 1 / (math.sqrt(i) + math.sqrt(i+1))

your_ans = round(your_ans, 2)
your_ans
```

Output:
```
24.0
```

```python
grader.check("q14")
```

Output:
```
q14 results: All test cases passed!
```

<!-- BEGIN QUESTION -->

### Question 15:

List all external sources (excluding the lecture notes and textbook), you used to answer the quiz questions, detailing the source and describing exactly what was learned and how that information was used in your assignment.

_Type your answer here, replacing this text._

<!-- END QUESTION -->



Tests for q1 (points: None)
```
>>> your_ans['dog'] == 4
True
```
```
>>> your_ans['cat'] == 2
True
```


Tests for q10 (points: None)
```
>>> your_ans == [13, 13, 19, 4, 6]
True
```


Tests for q12 (points: None)
```
>>> your_ans == 149823
True
```


Tests for q13 (points: None)
```
>>> your_ans == 'Bella'
True
```


Tests for q14 (points: None)
```
>>> your_ans == 24.0
True
```


Tests for q2 (points: None)
```
>>> highest == 91
True
```
```
>>> lowest == 12
True
```
```
>>> total_passed == 5
True
```


Tests for q3 (points: None)
```
>>> your_ans == 18
True
```


Tests for q4 (points: None)
```
>>> list(student_scores.index) == ['Alice', 'Bob', 'Charlie', 'David', 'Emma']
True
```
```
>>> list(student_scores.values) == [78, 92, 65, 88, 55]
True
```


Tests for q5 (points: None)
```
>>> fruits.shape == (5, 3)
True
```
```
>>> len(fruits.index.unique()) == 3
True
```


Tests for q6 (points: None)
```
>>> summer_fruits.shape == (2, 2)
True
```
```
>>> len(summer_fruits.columns) == 2
True
```


Tests for q7 (points: None)
```
>>> summer_fruits.shape == (2, 2)
True
```
```
>>> len(summer_fruits.columns) == 2
True
```


Tests for q8 (points: None)
```
>>> youngest == 'Justin'
True
```


Tests for q9 (points: None)
```
>>> your_ans == [15, 16, 24]
True
```
