```python
# Initialize Otter
import otter
grader = otter.Notebook("HW2.ipynb")
```

```python
import math
import pandas as pd
import numpy as np
```

**Due Date: Oct 5, 2026 by 11:59 pm**

Each question carries 5 points. 

# Homework Category

Each homework assignment contains **auxiliary-skill** and **advanced-skill** tasks.

The assignment mark is the **highest mark for which both skill thresholds are satisfied**.

## Homework Assignment Marks

| Assignment Mark |                                Auxiliary Tasks |                                 Advanced Tasks |
| --------------- | ---------------------------------------------: | ---------------------------------------------: |
| **E**           |                               At least **90%** |                               At least **80%** |
| **M**           |                               At least **80%** |                               At least **70%** |
| **N**           |                               At least **70%** |                               At least **60%** |
| **X**           | Below the N threshold or insufficient evidence | Below the N threshold or insufficient evidence |

> **Important:** Both skill requirements must be satisfied. Therefore, the **lower demonstrated skill determines the assignment mark**.

For example, a score of **95% on auxiliary tasks** and **65% on advanced tasks** results in an **N** mark.

---

## Homework Flexibility and Reassessment

### One-Week Extension

A student may extend the deadline of **one homework assignment by seven calendar days without penalty**.

The student must notify the instructor **by the original deadline**.

Approved absences and accommodations do **not** use this extension.

### One Revision

A student may revise **one homework assignment** within **seven calendar days after feedback is provided**.

The revision must include:

* an explanation of the original mistake,
* a description of the changes made, and
* evidence of improved reasoning.

The new evidence **replaces the previous mark**.

> Extensions and revisions may not extend beyond the published final submission deadline.

---

## Homework Category Grade

The final Homework Category grade is determined by the student's overall profile across the seven homework assignments.

| Category Grade | Required Homework Profile                                                                                 |
| -------------- | --------------------------------------------------------------------------------------------------------- |
| **A**          | **E** on at least **6 of the 7** homework assignments, with the remaining assignment at **N or higher**   |
| **B**          | **M or higher** on at least **6 of the 7** homework assignments; the remaining assignment cannot be **X** |
| **C**          | **N or higher** on at least **5 of the 7** homework assignments                                           |
| **D**          | **N or higher** on at least **4 of the 7** homework assignments                                           |
| **F**          | The student does not satisfy the **D** requirement                                                        |


## Instructions: 

1. After completing **HW2**, download your completed `HW2.ipynb` file by selecting **File → Download** in JupyterLab.

2. Upload your completed `HW2.ipynb` file to the **HW2 Dropbox in D2L**.

3. This assignment is due by the **date and time listed in the course calendar**.

4. This is an **individual assignment**. As stated in the course syllabus, the use of AI tools is prohibited. You may not receive assistance from anyone other than the instructor, the TA, or the AARC. Any help received from the TA or AARC must be clearly documented.

5. If you use any resource other than the instructor or the course textbook/notes, you must document it in a **new Markdown cell in Question 20 of this notebook**. Include the source, what you learned from it, and specifically how you used that information in completing the assignment. Submissions that are copied, either partially or entirely, from another source will be subject to a significant penalty.

6. If you need help, visit the instructor during the posted office hours. If those times do not work with your schedule, email the instructor to arrange an alternative meeting time.

7. Passing the visible autograder tests does **not** guarantee that your answer is correct. The instructor may use additional hidden tests during grading. Full credit will be awarded only when your answers also pass these tests.

8. If you encounter any technical or assignment-related difficulties, please contact the instructor for assistance.


# Auxiliary Skills

# 1. Numpy

<!-- BEGIN QUESTION -->

For each of the Python expressions below, write the output when the expression is evaluated. If an error occurs, write **Error**. Here is an example:

```python
np.array(1, 2, 3, 4, 5) == 2
```

Output: 
```python
array([False, True, False, False, False])
```

Note: Do NOT run this code before. First check your answer, and then you may run to verify your result. 

### Question 1(a)
```python
np.sum((np.arange(131, 138) - 135) >= 0)
```

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

### Question 1(b)
```python
np.array([8, 3, 4, 5]) + np.arange(2, 11, 3)
```

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

### Question 1(c)
```python
np.array([1, 2, 3]) + np.mean(np.arange(1, 10, 4))
```

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

### Question 1(d)
```python
np.max(np.arange(10))
```

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

### Question 1(e)
```python
temp = pd.DataFrame({'State': ['Alaska', 'Texas', 'California'], 'Avg Temp': [-7, 5, 8]})
np.mean(temp.iloc[:, 1])
```

_Type your answer here, replacing this text._

<!-- END QUESTION -->

# 2. Tables

A table named `instructor` contains one row for each SFA faculty member and these columns:
- **Dept**: a string, the department of the faculty member
- **Name**: a string, the first name fo the faculty member
- **Designation**: a string, one of: Lecturer, Assitant Professor, Associate Professor, or Professor
- **Salary**: an int, last year's salary paid by the university

Write a python code to answer the following questions: 

```python
instructor = pd.read_csv('data/instructor.csv')
instructor.head(3)
```

### Question 2(a)

The total salary amount paid to all faculty by the university last year. 

```python
your_ans = ...
your_ans
```

```python
grader.check("q2a")
```

### Question 2(b)

The name of the second highest paid faculty member. (Assume no two faculty have the same salary)

```python
your_ans = ...
your_ans
```

```python
grader.check("q2b")
```

### Question 2(c)

The average faculty salary after all instructors get a 10% raise in the present year. Round your answer upto two decimal points. 

```python
your_ans = ...
your_ans
```

```python
grader.check("q2c")
```

# Advance Skills

### Question 2(d)

Find the departments which have both Lecturers and Professors. Your answer should return a **numpy array** of departments.

You may use for loop to answer this question. Later, we will see better way (computationally fast) to solve it. 
Do not use concepts that are not covered in the lecture notes. 

```python

your_ans
```

```python
grader.check("q2d")
```

# 3. Gotham City

A numpy array of integers, named `gotham` contains the rough population of the Gotham city for every 10 years from 1900 to 2020. It has total 13 items. 

```
array([305500, 393300, ..., 25241897, 29217653])
```

Write a Python program to solve the following questions. Please note that since the full set of array elements is not provided, you will NOT BE able to execute the code to verify the results. 

<!-- BEGIN QUESTION -->

### Question 3 (a)

What is the population in the year 2000? 

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

### Question 3(b)

The population change from 1930 to 1950. You should only show the change, expressed as number of persons. (Not as proportion or percentage change). 

That is, population change from 1930 to 1950 will be the `population in 1950 - population in 1930`. 

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

### Question 3(c)

Find if the population ever **increased** by less than 500,000 in a decade. 

Note: Your answer should contain python expression which returns true or false. (Do not write True or False directly by yourself). Also, many students in the past have misread the question. **Be careful!**

Hint: Use `np.diff` method


_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

### Question 3(d)

Find the number of items in `gotham` that are atleast twice as large as the population in 1950. 

_Type your answer here, replacing this text._

<!-- END QUESTION -->

# Hurricane Harvey

Misinformation poses a significant challenge in today’s information era. This issue intensifies during natural
disasters, like hurricanes, when false information can spread rapidly. The table, harvey, below includes various tweets from
Hurricane Harvey in 2017, with some identified as false information or fake. The columns in the table are as follows:

- **Tweet**: _a string_, the tweet’s text that was posted by a twitter user
- **TweetID**: _a string_, a unique id for the posted tweet
- **ID**: _a string_, the user id who posted this tweet
- **Views**: _an int_, the number of unique users who viewed this tweet
- **isFake**: _a bool_, whether the tweet was fake or not.(True is fake)
- **Date**: _a string_, the month and day the tweet was posted

The first four rows of the table are shown below:

| Tweet | TweetID | ID | Views | isFake | Date |
| -- | -- | -- | -- | -- | -- |
| Scientists confirm that Hurricane Harvey was caused... | 01JA5WAMVQQ88 | A7B9C2D1 | 167692 | True | 08/28 | 
| National Weather Service has issued a flash flood...| 9SD8MD5JGTHHQ | Q1R2S3T6 | 160233 | False | 08/28 |
| BREAKING: Hurricane Harvey has changed course... | 5WAMWP66A88N6 | X3Y8Z5P4 | 1084894 | True | 09/01 |
| Local businesses are offering free boat rides... | 2P13QTKC1AJ56J7 | A7B9C2D1 | 520213 | True | 08/27|

Write python expressions to answer the following questions.

### Question 4(a)

Find the total number of fake tweets propagated during the hurricane.

```python
fake_tweets = harvey.query("...")
fake_tweets.shape...
```

### Question 4(b)

Find the date of the most viewed tweet? (Assume no two tweets have same number of views

```python
df = harvey.sort_values(...)
first_row = df.iloc...
first_row.loc...
```

### Question 4c

Find an array of all unique user ids who propagated misinformation during the hurricane.

Note: `fake_tweets` variable is already defined in part (a)

```python
np.unique(fake_tweets[...])
```

Tests for q2a (points: None)
```
>>> your_ans.item() == 53181784
True
```


Tests for q2b (points: None)
```
>>> your_ans == 'Sienna'
True
```


Tests for q2c (points: None)
```
>>> your_ans.item() == 116999.92
True
```


Tests for q2d (points: None)
```
>>> len(your_ans) == 19
True
```
