```python
# Initialize Otter
import otter
grader = otter.Notebook("Expressions.ipynb")
```

```python
import math
```

# Python as a Calculator

We can use python to act as a calculator. Python supports different numerical operators to perform such mathematical operations, such as: 

| Operator | Meaning | Example | Value | 
| -----   | ---- | ----- | ----- |
| * | Multiplication | 3 * 2 | 6 |
| / | Division | 3 / 2 | 1.5 |
| + | Addition | 3 + 2 | 5 |
| - | Subtraction | 3 - 2 | 1 |
| ** | Exponentiation | 3 ** 2 | 9 |
| % | Modulo | 3 % 2 | 1 |

All of the above operators are also called _binary operators_, as they need two data(binary) on each side of the operator (called operands) to compute the result or value. 

When a mathematical operation is made up of either an operator, or variable, or operand and produce an output, we call such an operation as `expression`. For eg: 

```python
5 + 2 * 4
```

```python
x = 5
x * 3
```

### Question 1

Write a python expression to calculate the value of $4^5$. Store your result in a variable `your_ans`. 

```python
your_ans = ...
your_ans
```

```python
grader.check("q1")
```

### Question 2

What is the output of the python expression `17 % 4`? Do not write code for this, simply store your answer in the variable `your_ans` below in a markdown cell.

Tests for q1 (points: None)
```
>>> your_ans == 1024
True
```
