```python
# Initialize Otter
import otter
grader = otter.Notebook("Numpy_Part3.ipynb")
```

```python
import numpy as np
```

# Numpy Part3

In Part 2, we learned how to perform computations between a NumPy array and a scalar, as well as element-wise computations between arrays. We also explored several NumPy functions used for mathematical calculations.
In this lesson, we will study two especially useful NumPy features:
- **Boolean masking**, which lets us filter data using conditions.
- **Broadcasting**, which lets NumPy perform calculations on arrays with compatible, but different shapes.

These features help us replace many manual loops with code that is shorter, clearer, and usually faster. Let's start!

A **Boolean array** is a NumPy array whose elements are either `True` or `False`. We can create one directly, but it is more common to create one by comparing numerical data.

```python
barr1 = np.array([True, True, False, False, True])
barr2 = np.array([True, False, False, True, True])

print("barr1:", barr1)
print("barr2:", barr2)
print("dtype:", barr1.dtype)
```

## Comparisons create Boolean arrays

When we compare an array with a number, NumPy performs the comparison element by element.

```python
temperatures = np.array([72, 81, 95, 68, 88])

hot_days = temperatures >= 85

print("Temperatures:", temperatures)
print("At least 85°F:", hot_days)
```

## Combining Boolean conditions

For NumPy arrays, use these element-wise logical operators:

| Meaning | Operator | Function equivalent |
|---|:---:|---|
| AND | `&` | `np.logical_and(...)` |
| OR | `\|` | `np.logical_or(...)` |
| NOT | `~` | `np.logical_not(...)` |

The operators compare values at matching positions.

```python
barr1
```

```python
barr2
```

```python
print("AND:", barr1 & barr2)
print(" OR:", barr1 | barr2)
print("NOT barr1:", ~barr1)

# NumPy has equivalent named functions which produce the same results.
print("logical_and:", np.logical_and(barr1, barr2))
print("logical_or:", np.logical_or(barr1, barr2))
print("logical_not:", np.logical_not(barr1))
```

### Boolean Masking

A Boolean mask is a Boolean array used to select elements from another array.
Think of a mask as a row of switches:
- **True** means keep this value.
- **False** means leave this value out.

The expression ```data[mask]``` returns the values wherever the mask is `True`.

For example: 

<img src="pics/bm1.drawio.png" width=700 />

```python
arr = np.array([10, 20, 30, 40])
mask = np.array([True, True, False, True])
arr[mask]
```

For example: 
    
<img src="pics/bm2.drawio.png" width=700 />

```python
arr = np.array([[10, -20, 30], [-1, 3, 5], [-6, -4, 9]])
arr
```

```python
mask = arr > 0
mask
```

```python
arr[mask]
```

Usually, we do not type the mask manually. We create it from a condition and use it to filter the original array.

```python
temperatures = np.array([72, 81, 95, 68, 88])
hot_days = temperatures >= 85

print("Hot temperatures:", temperatures[hot_days])

# The condition can also be written directly inside the brackets.
print("Same result:", temperatures[temperatures >= 85])
```

## Filtering with multiple conditions

Suppose we want temperatures from 70°F through 89°F, inclusive. Both conditions must be true, so we use &.

```python
mask1 = temperatures >= 70
mask2 = temperatures < 90
comfortable_mask = mask1 & mask2
## comfortable_mask = np.logical_and(mask1, mask2)
comfortable_temperatures = temperatures[comfortable_mask]

print("Mask:", comfortable_mask)
print("Comfortable temperatures:", comfortable_temperatures)
```

## Choosing between two values with np.where()
`np.where(condition, value_if_true, value_if_false)` makes an element-wise choice. It is useful when we want to label or transform every element rather than remove some elements.

```python
scores = np.array([82, 59, 91, 76, 64])
results = np.where(scores >= 70, "Pass", "Review")

print(results)
```

# Numpy Broadcasting

**Broadcasting** is NumPy's ability to perform element-wise operations on arrays with different but compatible shapes.
Instead of physically copying smaller data many times, NumPy behaves as if the smaller array were stretched to fit the larger one. This makes calculations concise and memory-efficient.

### Example 1: A scalar with an array

A scalar has no axes, so NumPy can broadcast it to every element of an array.

```python
prices = np.array([10.00, 15.00, 22.50, 8.00])
sale_prices = prices * 0.80

print("Original prices:", prices)
print("20% off:        ", sale_prices)
```

### Broadcasting rule

NumPy compares shapes from `right` to `left`. Two dimensions are compatible when:
1. They are equal, **or**
2. One of them is `1`.
If one shape has fewer dimensions, its missing dimensions are treated as 1 on the left.

| Shape A | Shape B | Compatible? | Broadcast result |
|---|---|:---:|---|
| `(3, 4)` | `(4,)` | Yes | `(3, 4)` |
| `(3, 4)` | `(1, 4)` | Yes | `(3, 4)` |
| `(3, 4)` | `(3, 1)` | Yes | `(3, 4)` |
| `(3, 4)` | `(3,)` | No | — |
| `(2, 3)` | `(3, 2)` | No | — |

For example, compare (3, 4) with (4,) by aligning dimensions on the right:

(3, 4)
(   4)

### Example 2: Broadcasting a row vector
The matrix has shape (3, 4), and the vector has shape (4,). NumPy matches the four vector values with the four columns and applies the vector to every row.

The last dimensions match, and the missing dimension behaves like 1, so broadcasting succeeds.

```python
sales = np.array([
    [10, 20, 30, 40],
    [15, 25, 35, 45],
    [12, 22, 32, 42]
])
column_adjustments = np.array([1, 2, 3, 4])

adjusted_sales = sales + column_adjustments

print("sales shape:", sales.shape)
print("adjustments shape:", column_adjustments.shape)
print("\nAdjusted sales:\n", adjusted_sales)
```

The operation behaves conceptually as if NumPy used this larger array:
```text
[[1, 2, 3, 4],
 [1, 2, 3, 4],
 [1, 2, 3, 4]]  
```

NumPy does not need us to create that repeated array manually.

## Combining broadcasting and Boolean masking
These features are especially powerful together. Broadcasting can compare an entire table with a vector of requirements, producing a Boolean mask that we can count or use for filtering.


### Worked example: Exam performance
Four students completed three exams. Each exam receives a different bonus. We will:
1. Add the bonus to every student's scores using broadcasting.
2. Create a passing mask.
3. Count passed exams for each student.
4. Select students whose adjusted average is at least 80.

```python
student_names = np.array(["Ava", "Ben", "Camila", "Dev"])

exam_scores = np.array([
    [78, 84, 91],
    [65, 72, 68],
    [88, 79, 95],
    [70, 76, 82]
])

exam_bonus = np.array([2, 1, 3])  # one bonus for each exam
adjusted_scores = exam_scores + exam_bonus

print("Score shape:", exam_scores.shape)
print("Bonus shape:", exam_bonus.shape)
print("\nAdjusted scores:\n", adjusted_scores)
```

```python
passing_mask = adjusted_scores >= 70
passed_exam_counts = np.sum(passing_mask, axis=1)
student_averages = np.mean(adjusted_scores, axis=1)
high_average_mask = student_averages >= 80

print("Passing mask:\n", passing_mask)
print("\nNumber of exams passed:", passed_exam_counts)
print("Adjusted averages:", student_averages)
print("Students averaging at least 80:", student_names[high_average_mask])
```

### Understanding axis

For a 2D array:

| Expression | Direction collapsed | Result describes |
|---|---|---|
| `array.sum(axis=0)` | Rows | One result per column |
| `array.sum(axis=1)` | Columns | One result per row |

In the previous example, each row belongs to one student. Therefore, `np.sum(passing_mask, axis=1)` counts the `True` values across that student's exams.

# KC4_Part3

Answer the following questions and upload this notebook to **KC4_Part3 dropbox**. The due date is **Sep 24 by 11:59 pm.**

Do not use `for` loops to answer these questions. Try to use numpy concepts to solve them.

### Question 1

Write a python code to find the sum of all positive integers in the matrix $A= \begin{bmatrix}
1 & -2 & -3 & 4 \\
1 & 1 & -1 & -1 \\
2 & -2 & -1 & 0 \\
\end{bmatrix}$

```python
A = np.array([[1, -2, -3, 4], [1, 1, -1, -1], [2, -2, -1, 0]]) 
your_ans = ...
your_ans
```

```python
grader.check("q1")
```

### Question 2

Write a python code to find the sum of the positive integers present in the second row of the matrix A (defined in Question 1). 

```python
your_ans = ...
your_ans
```

```python
grader.check("q2")
```

### Question 3

Write a python code to count the total number of zeros present in the matrix 

$$
A = \begin{bmatrix}
1 & 5 & 0 & 4 \\
1 & 3 & 0 & 0 \\
2 & -2 & -1 & 0 \\
4 & 0 & 9  & 0  \\
2 & 9 & 0 & 0 \\
\end{bmatrix}
$$

```python
A = np.array([[1, 5, 0, 4], [1, 3, 0, 0], [2, -2, -1, 0], [4, 0, 9, 0], [2, 9, 0, 0]])
your_ans = ...
your_ans
```

```python
grader.check("q3")
```

### Question 4

Write a python code to replace all the negative numbers with 0 in the given array $A= \begin{bmatrix}
1 & -4 & -1 & 4 \\
1 & 3 & -8 & 1 \\
\end{bmatrix}$. 

Your resultant array A will be $\begin{bmatrix}
1 & 0 & 0 & 4 \\
1 & 3 & 0 & 1 \\
\end{bmatrix}$. 

```python
A = np.array([[1, -4, -1, 4], [1, 3, -8, 1]])
mask = ...
A[mask] = ...
A
```

```python
grader.check("q4")
```

### Question 5

Select the applicants who have at least 2 years of experience and a test score of at least 80.


```python
applicants = np.array(["Alex", "Bob", "Charlie", "David", "Emily"])
experience = np.array([1, 4, 2, 5, 3])
test_scores = np.array([92, 76, 85, 88, 79])
```

```python
qualified_applicants = ...

print("Qualified applicants:", qualified_applicants)
```

```python
grader.check("q5")
```

Tests for q1 (points: None)
```
>>> your_ans.item() == 9
True
```


Tests for q2 (points: None)
```
>>> your_ans.item() == 2
True
```


Tests for q3 (points: None)
```
>>> your_ans.item() == 8
True
```


Tests for q4 (points: None)
```
>>> A.shape == (2, 4)
True
```
```
>>> A[0, 2].item() == 0 and A[0, 1].item() == 0 and (A[1, 2].item() == 0)
True
```


Tests for q5 (points: None)
```
>>> list(qualified_applicants) == ['Charlie', 'David']
True
```
