```python
# Initialize Otter
import otter
grader = otter.Notebook("Lab3.ipynb")
```

```python
import math
import pandas as pd
import numpy as np
```

Before working on the Lab, make sure you **run the Python code cells at the top of the notebook**. These cells initialize the autograder and are required for the tests to work correctly.

* If the first Python cell produces an error, try **restarting the kernel** by selecting **Kernel → Restart Kernel** from the JupyterLab menu.
* After restarting the kernel, run the first Python cell again, followed by the remaining setup cells.
* Make sure these setup cells run successfully before attempting the autograder tests.

### Due Dates

**Core Skills:** Core Skill questions are due **during class**. If you are unable to complete all of them in class, you may submit the remaining questions with your Auxiliary Skill work. However, any Core Skill question completed after class can earn at most **50% of its available points**, even if the answer is correct.

**Auxiliary Skills:** Auxiliary Skill questions are due **Oct 2 by 11:59 p.m.**

## Laboratory Grading Policy

Each lab includes two types of tasks:

* **Core Tasks** – demonstrate your understanding of the essential skills and must be completed during the scheduled lab period.
* **Auxiliary Tasks** – build on the core skills and may usually be completed after class.

Each question carries **5 points**.

### Lab Marks

| Mark  |                                       Core Tasks |                                  Auxiliary Tasks |
| ----- | -----------------------------------------------: | -----------------------------------------------: |
| **E** |                                     At least 90% |                                     At least 90% |
| **M** |                                     At least 80% |                                     At least 70% |
| **N** |                                     At least 70% |                                     At least 60% |
| **X** | Below the N requirement or insufficient evidence | Below the N requirement or insufficient evidence |

### Overall Laboratory Grade

There are **six laboratory assignments**.

| Course Grade | Required Lab Performance                                            |
| ------------ | ------------------------------------------------------------------- |
| **A**        | **E** on at least 5 labs; the remaining lab must be **N or higher** |
| **B**        | **M or higher** on at least 5 labs; the remaining lab may be **X**  |
| **C**        | **N or higher** on at least 5 labs                                  |
| **D**        | **N or higher** on at least 4 labs                                  |
| **F**        | Does not meet the requirements for a D                              |

### Lab Completion Rules

**Approved Absences and Accommodations:** The 50% Core Task limit does not apply to approved institutional absences, instructor-approved excused absences, or applicable accommodations. A make-up lab or practical demonstration will be provided when appropriate.

**One Lab Reassessment:** You may reassess **one lab marked N or X**. The reassessment must be completed within **7 calendar days after feedback is provided** and may include a corrected submission, a practical demonstration, or both. The revised mark may be **M or E** if you independently demonstrate the required level of proficiency.


## Instructions: 

1. After completing **Lab3**, download your completed `Lab3.ipynb` file by selecting **File → Download** in JupyterLab.

2. Upload your completed `Lab3.ipynb` file to the **Lab3 Dropbox in D2L**.

3. This assignment is due by the **date and time listed in the course calendar**.

4. This is an **individual assignment**. As stated in the course syllabus, the use of AI tools is prohibited. You may not receive assistance from anyone other than the instructor, the TA, or the AARC. Any help received from the TA or AARC must be clearly documented.

5. If you use any resource other than the instructor or the course textbook/notes, you must document it in a **new Markdown cell after the last question of this notebook**. Include the source, what you learned from it, and specifically how you used that information in completing the assignment. Submissions that are copied, either partially or entirely, from another source will be subject to a significant penalty.

6. If you need help, visit the instructor during the posted office hours. If those times do not work with your schedule, email the instructor to arrange an alternative meeting time.

7. Passing the visible autograder tests does **not** guarantee that your answer is correct. The instructor may use additional hidden tests during grading. Full credit will be awarded only when your answers also pass these tests.

8. If you encounter any technical or assignment-related difficulties, please contact the instructor for assistance.


# Core Skills (Due in-class)

## A. Numpy

### Question 1:

Create a numpy array containing the first 100 squares i.e, $1^2, 2^2, 3^2, 4^2, 5^2, 6^2, ....., 100^2$  = 1, 4, 9, 16, 25, 36, 49, 64, ...., 10000 and add them up. Store your answer in the variable `sum_of_squares`

```python
arr = ...
sum_of_squares = ...
sum_of_squares
```

```python
grader.check("q1")
```

### Question 2:

You are given the following matrix, $A = \begin{bmatrix}
1 & 2 & 3 & 4 \\
10 & 20 & 30 & 40 \\
100 & 200 & 300 & 400 \\
-10 & -20 & -30 & -40 \\
-1 & -2 & -3 & -4 \\
\end{bmatrix}$


Your task is to write a python code to add all the elements of 3rd row (100, 200, 300, 400) from the matrix A. Store your answer in the variable `your_ans`

```python
A = np.array([
             [1, 2, 3, 4],
             [10, 20, 30, 40],
             [100, 200, 300, 400],
             [-10, -20, -30, -40],
             [-1, -2, -3, -4]
            ])
A
```

```python
your_ans = ...
your_ans
```

```python
grader.check("q2")
```

### Question 3:

You are given a matrix `A` as in the previous question.


Your task is to write a python code to add all the elements of the 2nd and 3rd row and 1st and 4th column from the matrix A. 
```python
10 40
100 400
```

Store your answer in the variable `your_ans`

Note: You must not add these numbers manually. Your code should work for any other values of the matrix A.

```python
your_ans = ...
your_ans 
```

```python
grader.check("q3")
```

### Question 4:

You are given the following matrix, $B = \begin{bmatrix}
\hspace{0.4cm}1 & -2 & \hspace{0.4cm}3 & \hspace{0.4cm}4 \\
\hspace{0.4cm}1 & \hspace{0.4cm}2 & -3 & \hspace{0.4cm}4 \\
\hspace{0.4cm}1 & \hspace{0.4cm}2 & \hspace{0.4cm}3 & \hspace{0.4cm}4 \\
-1 & \hspace{0.4cm}2 & \hspace{0.4cm}3 & -4 \\
-1 & -2 & \hspace{0.4cm}3 & \hspace{0.4cm}4 \\
\end{bmatrix}$


Your task is to write a python code to add all the negative numbers of the matrix A. 

Store your answer in the variable `your_ans`

Note: You must not add these numbers manually. Your code should work for any other values of the matrix A.

**Hint**: Use masking

```python
B = np.array([
             [1, -2, 3, 4],
             [1, 2, -3, 4],
             [1, 2, 3, 4],
             [-1, 2, 3, -4],
             [-1, -2, 3, 4]
            ])
B
```

```python
mask = ...
your_ans = np.sum(...) # Apply the mask
your_ans
```

```python
grader.check("q4")
```

### Question 5:

Suppose you are given two matrices A and B, where  $A = \begin{bmatrix}
1 & 2 \\
1 & 1 \\
10 & 20 \\
\end{bmatrix}$ and $B = \begin{bmatrix}
1 & 1 \\
2 & 2 \\
3 & 3 \\
\end{bmatrix}$ . Write a python code to add these two matrices together and store your answer in a variable, `C`.

```python
A = np.array([[1, 2], [1, 1], [10, 20]])
B = np.array([[1, 1], [2, 2], [3, 3]])
C = ...
C
```

```python
grader.check("q5")
```


# B. D2L Assignment Submission 

**D2L Dropbox** is a feature of D2L that allows students to submit assignments for their courses. Students can submit assignments until the specified due date. Submissions made after the deadline are accepted but may incur a late penalty.

For this assignment, assume the following:

- Each student can submit an assignment only once.
- Instructors can provide feedback after an assignment has been submitted.
- An instructor may teach multiple courses during a semester.
- A student may enroll in multiple courses taught by the same instructor.

The DataFrame `dropbox` contains one row for each assignment submission made by SFA students during Spring 2026.

### Description of the DataFrame

| Column | Data Type | Description |
|---|---|---|
| `username` | String | Username of the student |
| `title` | String | Title of the assignment |
| `date` | String | Date on which the assignment was submitted |
| `level` | String | Undergraduate classification (freshman, sophomore, junior, or senior) |
| `late` | Float | Number of minutes the assignment was submitted late |

The first three rows of the DataFrame are shown below.

| username | title | date | level | late |
|---|---|---|---|---:|
| davis29 | csci4335_homework1 | 3/23 | senior | 0.0 |
| samf1 | csci1462_homework1 | 3/12 | freshman | 90.2 |
| vicky | csci1462_homework1 | 4/5 | sophomore | 35.5 |

**Note:** An assignment is considered late when its `late` value is greater than zero. Hints are provided to guide you, but you may use alternative pandas techniques to solve these problems.

<!-- BEGIN QUESTION -->

### Question 6  

Find the total number of students enrolled in the Spring 2025. 

**Hint:** Use `np.unique()` to determine the unique students enrolled.

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

### Question 7

Find the **total number of assignments submitted late** by the student `davis29`.

**Hint:** Use `query()` or Boolean masking to identify the relevant submissions.

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

### Question 8

Find the **total number of late assignment submissions for each undergraduate level** (freshman, sophomore, junior, and senior).

**Hint:** Filter the late submissions and use `groupby()` to summarize the results.

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

### Question 9

Find the **total number of assignments submitted late by `davis29` during March**.

**Hint:** You may use `Series.str.startswith('3/')` to create a Boolean mask identifying submissions made in March. Combine it with the appropriate conditions to identify the required submissions.

_Type your answer here, replacing this text._

<!-- END QUESTION -->

# Auxiliary Skills

<!-- BEGIN QUESTION -->

### Question 10

Find the number of students who submitted **all their assignments on time** during Spring 2025.

Hint: A student should be included only if none of their assignment submissions were late. Even if a student submitted just one assignment late, exclude that student from your final answer.

_Type your answer here, replacing this text._

<!-- END QUESTION -->

### Question 11:

#### Estimate $\pi$ using Walli's Formula

About half a centure before Leibniz, John Wallis gave an expression to estimate the value of $\pi$ as an infinite product. 

$$\pi = 2 \cdot \left(\frac{2}{1} \cdot \frac{2}{3} \cdot \frac{4}{3} \cdot \frac{4}{5} \cdot \frac{6}{5} \cdot \frac{6}{7} \cdots \right) $$

This is a product of "even/odd" fractions. Let's use arrays to multiple a million of them, and see if the product is close to $\pi$. 

$$\pi = 2 \cdot \left(\frac{2}{1} \cdot \frac{4}{3} \cdot \frac{6}{5} \cdots \frac{1000000}{999999} \right) \cdot \left(\frac{2}{3} \cdot \frac{4}{5} \cdot \frac{6}{7} \cdots \frac{1000000}{1000001} \right) $$

We start by creating an array of even numbers 2, 4, 6, and so on upto 1000000; and odd numbers: 1, 3, 5, 7, ... upto 99999 and 3, 5, 7, ... upto 1000001. 

**Hint**: Look into `np.prod()` method

```python
even_numbers = ...
odd_numbers = ...

left_eq = ...
right_eq = ...

pi = 2 * left_eq * right_eq
pi
```

```python
grader.check("q11")
```

## C. Nvidia stock

A numpy array of integers, named `nvidia` contains the rough adjusted closing monthly stock prices (in dollars) of NVIDIA corporation, a key innovator in computer graphics, and AI technology, for the year 2023. It has total 12 items.

```python
nvidia = np.array([19.52, 23.20, 27.76, 27.74, 37.82, 42.28, 46.71, 49.34, 43.48, 40.77, 46.76, 49.51])
nvidia
```

Write a Python code to answer each of the following questions.

#### Question 12
What is the stock price value in October? Round your answer upto two decimal points.

```python
your_ans = ...
your_ans
```

```python
grader.check("q12")
```

### Question 13

The stock price change for each consecutive month from Jan to Dec. For eg: the change of stock price between month of Feb and Jan is 23.2 - 19.52 = 3.68. Round your answer upto two decimal points. Your final answer should be a numpy array with 11 elements. Store your answer in the variable, `nvidia_diff`. 

```python
nvidia_diff = ...
nvidia_diff
```

```python
grader.check("q13")
```

### Question 14

Check if the stock prices ever rose by more than 4 dollars between consecutive months. Your solution should include Python code that returns either True or False (without explicitly writing True or False in the answer).

```python
nvidia_diff # use this variable to answer the question
```

```python
your_ans = ...
your_ans
```

```python
grader.check("q14")
```

### Question 15

Determine how many times the stock prices increased by more than 4 dollars between consecutive months. Your answer should include Python code that returns either True or False (without directly writing True or False in your answer).

```python
your_ans = ...
your_ans
```

```python
grader.check("q15")
```

### Question 16:

Determine how many times the stock prices ever **rose** by less than 4 dollars between consecutive months. Your solution should include Python code that returns either True or False (without explicitly writing True or False in the answer).

```python
nvidia_diff # Use this variable to answer the question
```

```python
your_ans = ...
your_ans
```

```python
grader.check("q16")
```

## D. Hurricanes:

The `hurricanes` dataset contains information on hurricanes that affect the continent of North America. The columns are as follows:

**Year** -*an int*,  the year when hurricane hit

**Name** - *a string*, the name of the hurricane

**Rain Inch** - *a float*,  the amount of rain that fell in inches

**Highest Wind Speed** - *an int*, The highest wind speed achieved by the hurricane

**Damage(USD)** - *a float*, the cost of damage in US dollars

**Fatalities** - *an int*, the amount of deaths

**Areas Affected** - *a string*, the area affected by the hurricane

```python
hurricanes = pd.read_csv('data/hurricanes.csv')
hurricanes.head(3)
```

### Question 17

Find the hurricane name that resulted in the highest number of fatalities in a particular year. Assume there is only one.

```python
your_ans = ...
your_ans
```

```python
grader.check("q17")
```

### Question 18:

Create a new dataframe called `hurricane_normalized` that includes the `name` column and a `norm_rain` column, where `name` represents the hurricane's name and `norm_rain` is the normalized amount of rainfall in inches.

To normalize a column, use the formula:

$$z = \frac{x-\mu}{\sigma}$$

where,
- $x$ is the value from the column,
- $\mu$ is the mean of the values in that column,
- $\sigma$ is the standard deviation of the values in that column, 
- $z$ is the normalized value for that column.

```python
mu_rain = np.mean(hurricanes['Rain Inch.']) # calculates the mean of the column
sigma_rain = np.std(hurricanes['Rain Inch.']) # calcualtes the standard deviation of the column 

hurricanes_normalized = pd.DataFrame({ 
                            ...
                        })


hurricanes_normalized.head(3)
```

```python
grader.check("q18")
```

### Question 19

Find the hurricane name that occured most frequently in the recorded dataset. Assume there is only one such hurricane.

```python
your_ans = ...
your_ans
```

```python
grader.check("q19")
```

### Question 20

Find the names of different hurricanes that hit `Puerto Rico` throughout the years. Store your answer in a variable `your_ans` as a numpy array. 

**Hint**: Revisit section `String operation on Series` in _Pandas_Series_DataFrame_Tutorial_ notebook. 
Look into `np.unique` method to find the names of different hurricanes

```python
your_ans = ...
your_ans
```

```python
grader.check("q20")
```

Tests for q1 (points: None)
```
>>> sum_of_squares.item() == 338350
True
```


Tests for q11 (points: None)
```
>>> round(pi, 2).item() == 3.14
True
```


Tests for q12 (points: None)
```
>>> np.isclose(40.77, your_ans).item()
True
```


Tests for q13 (points: None)
```
>>> len(nvidia_diff) == 11
True
```
```
>>> np.isclose(nvidia_diff[-1], 2.75).item()
True
```


Tests for q14 (points: None)
```
>>> your_ans.item() == True
True
```


Tests for q15 (points: None)
```
>>> your_ans == 5
True
```


Tests for q16 (points: None)
```
>>> your_ans == 3
True
```


Tests for q17 (points: None)
```
>>> your_ans == 'Maria'
True
```


Tests for q18 (points: None)
```
>>> hurricanes_normalized.shape == (124, 2)
False
```


Tests for q19 (points: None)
```
>>> your_ans == 'Barry'
True
```


Tests for q2 (points: None)
```
>>> your_ans.item() == 1000
True
```


Tests for q20 (points: None)
```
>>> len(your_ans) == 17
True
```


Tests for q3 (points: None)
```
>>> your_ans.item() == 550
True
```


Tests for q4 (points: None)
```
>>> your_ans.item() == -13
True
```


Tests for q5 (points: None)
```
>>> np.sum(C).item() == 47
True
```
```
>>> C.shape == (3, 2)
True
```
